Functionally-Coupled Ion Channels Begin Co-assembling at the Start of Their Synthesis

Author response:

The following is the authors’ response to the original reviews.

Recommendations for the Authors:

(1) Clarify Mechanistic Interpretations

(a) Provide stronger evidence or a more cautious interpretation regarding whether intracellular BK-CaV1.3 ensembles are precursors to plasma membrane complexes.

This is an important point. We adjusted the interpretation regarding intracellular BKCaV1.3 hetero-clusters as precursors to plasma membrane complexes to reflect a more cautious stance, acknowledging the limitations of available data. We added the following to the manuscript.

“Our findings suggest that BK and CaV1.3 channels begin assembling intracellularly before reaching the plasma membrane, shaping their spatial organization and potentially facilitating functional coupling. While this suggests a coordinated process that may contribute to functional coupling, further investigation is needed to determine the extent to which these hetero-clusters persist upon membrane insertion.”

(b) Discuss the limitations of current data in establishing the proportion of intracellular complexes that persist on the cell surface.

We appreciate the suggestion. We expanded the discussion to address the limitations of current data in determining the proportion of intracellular complexes that persist on the cell surface. We added the following to the manuscript.

“Our findings highlight the intracellular assembly of BK-CaV1.3 hetero-clusters, though limitations in resolution and organelle-specific analysis prevent precise quantification of the proportion of intracellular complexes that ultimately persist on the cell surface. While our data confirms that hetero-clusters form before reaching the plasma membrane, it remains unclear whether all intracellular hetero-clusters transition intact to the membrane or undergo rearrangement or disassembly upon insertion. Future studies utilizing live cell tracking and high resolution imaging will be valuable in elucidating the fate and stability of these complexes after membrane insertion.”

(2) Refine mRNA Co-localization Analysis

(a) Include appropriate controls using additional transmembrane mRNAs to better assess the specificity of BK and CaV1.3 mRNA co-localization.

We agree with the reviewers that these controls are essential. We explain better the controls used to address this concern. We added the following to the manuscript.

“To explore the origins of the initial association, we hypothesized that the two proteins are translated near each other, which could be detected as the colocalization of their mRNAs (Figure 5A and B). The experiment was designed to detect single mRNA molecules from INS-1 cells in culture. We performed multiplex in situ hybridization experiments using an RNAScope fluorescence detection kit to be able to image three mRNAs simultaneously in the same cell and acquired the images in a confocal microscope with high resolution. To rigorously assess the specificity of this potential mRNA-level organization, we used multiple internal controls. GAPDH mRNA, a highly expressed housekeeping gene with no known spatial coordination with channel mRNAs, served as a baseline control for nonspecific colocalization due to transcript abundance. To evaluate whether the spatial proximity between BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) was unique to functionally coupled channels, we also tested for NaV1.7 mRNA (SCN9A), a transmembrane sodium channel expressed in INS-1 cells but not functionally associated with BK. This allowed us to determine whether the observed colocalization reflected a specific biological relationship rather than shared expression context. Finally, to test whether this proximity might extend to other calcium sources relevant to BK activation, we probed the mRNA of ryanodine receptor 2 (RyR2), another Ca2+ channel known to interact structurally with BK channels [32]. Together, these controls were chosen to distinguish specific mRNA colocalization patterns from random spatial proximity, shared subcellular distribution, or gene expression level artifacts.”

(b) Quantify mRNA co-localization in both directions (e.g., BK with CaV1.3 and vice versa) and account for differences in expression levels.

We thank the reviewer for this suggestion. We chose to quantify mRNA co-localization in the direction most relevant to the formation of functionally coupled hetero-clusters, namely, the proximity of BK (KCNMA1) mRNA to CaV1.3 (CACNA1D) mRNA. Since BK channel activation depends on calcium influx provided by nearby CaV1.3 channels, this directional analysis more directly informs the hypothesis of spatially coordinated translation and channel assembly. To address potential confounding effects of transcript abundance, we implemented a scrambled control approach in which the spatial coordinates of KCNMA1 mRNAs were randomized while preserving transcript count. This control resulted in significantly lower colocalization with CACNA1D mRNA, indicating that the observed proximity reflects a specific spatial association rather than expressiondriven overlap. We also assessed colocalization of CACNA1D with both KCNMA1, GAPDH mRNAs and SCN9 (NaV1.7); as you can see in the graph below these data support t the same conclusion but were not included in the manuscript.

(c) Consider using ER labeling as a spatial reference when analyzing mRNA localization

We thank the reviewers for this suggestion. Rather than using ER labeling as a spatial reference, we assess BK and CaV1.3 mRNA localization using fluorescence in situ hybridization (smFISH) alongside BK protein immunostaining. This approach directly identifies BK-associated translation sites, ensuring that observed mRNA localization corresponds to active BK synthesis rather than general ER association. By evaluating BK protein alongside its mRNA, we provide a more functionally relevant measure of spatial organization, allowing us to assess whether BK is synthesized in proximity to CaV1.3 mRNA within micro-translational complexes. The results added to the manuscript is as follows.

“To further investigate whether KCNMA1 and CACNA1D are localized in regions of active translation (Figure 7A), we performed RNAScope targeting KCNMA1 and CACNA1D alongside immunostaining for BK protein. This strategy enabled us to visualize transcript-protein colocalization in INS-1 cells with subcellular resolution. By directly evaluating sites of active BK translation, we aimed to determine whether newly synthesized BK protein colocalized with CACNA1D mRNA signals (Figure 7A). Confocal imaging revealed distinct micro-translational complex where KCNMA1 mRNA puncta overlapped with BK protein signals and were located adjacent to CACNA1D mRNA (Figure 7B). Quantitative analysis showed that 71 ± 3% of all KCNMA1 colocalized with BK protein signal which means that they are in active translation. Interestingly, 69 ± 3% of the KCNMA1 in active translation colocalized with CACNA1D (Figure 7C), supporting the existence of functional micro-translational complexes between BK and CaV1.3 channels.”

(3) Improve Terminology and Definitions

(a) Clarify and consistently use terms like "ensemble," "cluster," and "complex," especially in quantitative analyses.

We agree with the reviewers, and we clarified terminology such as 'ensemble,' 'cluster,' and 'complex' and used them consistently throughout the manuscript, particularly in quantitative analyses, to enhance precision and avoid ambiguity.

(b) Consider adopting standard nomenclature (e.g., "hetero-clusters") to avoid ambiguity.

We agree with the reviewers, and we adapted standard nomenclature, such as 'heteroclusters,' in the manuscript to improve clarity and reduce ambiguity.

(4) Enhance Quantitative and Image Analysis

(a) Clearly describe how colocalization and clustering were measured in super-resolution data.

We thank the reviewers for this suggestion. We have modified the Methods section to provide a clearer description of how colocalization and clustering were measured in our super-resolution data. Specifically, we now detail the image processing steps, including binary conversion, channel multiplication for colocalization assessment, and density-based segmentation for clustering analysis. These updates ensure transparency in our approach and improve accessibility for readers, and we added the following to the manuscript.

“Super-resolution imaging:

Direct stochastic optical reconstruction microscopy (dSTORM) images of BK and 1.3 overexpressed in tsA-201 cells were acquired using an ONI Nanoimager microscope equipped with a 100X oil immersion objective (1.4 NA), an XYZ closed-loop piezo 736 stage, and triple emission channels split at 488, 555, and 640 nm. Samples were imaged at 35°C. For singlemolecule localization microscopy, fixed and stained cells were imaged in GLOX imaging buffer containing 10 mM β-mercaptoethylamine (MEA), 0.56 mg/ml glucose oxidase, 34 μg/ml catalase, and 10% w/v glucose in Tris-HCl buffer. Single-molecule localizations were filtered using NImOS software (v.1.18.3, ONI). Localization maps were exported as TIFF images with a pixel size of 5 nm. Maps were further processed in ImageJ (NIH) by thresholding and binarization to isolate labeled structures. To assess colocalization between the signal from two proteins, binary images were multiplied. Particles smaller than 400 nm2 were excluded from the analysis to reflect the spatial resolution limit of STORM imaging (20 nm) and the average size of BK channels. To examine spatial localization preference, binary images of BK were progressively dilated to 20 nm, 40 nm, 60 nm, 80 nm, 100 nm, and 200 nm to expand their spatial representation. These modified images were then multiplied with the CaV1.3 channel to quantify colocalization and determine BK occupancy at increasing distances from CaV1.3. To ensure consistent comparisons across distance thresholds, data were normalized using the 200 nm measurement as the highest reference value, set to 1.”

(b) Where appropriate, quantify the proportion of total channels involved in ensembles within each compartment.

We thank the reviewers for this comment. However, our method does not allow for direct quantification of the total number of BK and CaV1.3 channels expressed within the ER or ER exit sites, as we rely on proximity-based detection rather than absolute fluorescence intensity measurements of individual channels. Traditional methods for counting total channel populations, such as immunostaining or single-molecule tracking, are not applicable to our approach due to the hetero-clusters formation process. Instead, we focused on the relative proportion of BK and CaV1.3 hetero-clusters within these compartments, as this provides meaningful insights into trafficking dynamics and spatial organization. By assessing where hetero-cluster preferentially localize rather than attempting to count total channel numbers, we can infer whether their assembly occurs before plasma membrane insertion. While this approach does not yield absolute quantification of ER-localized BK and CaV1.3 channels, it remains a robust method for investigating hetero-cluster formation and intracellular trafficking pathways. To reflect this limitation, we added the following to the manuscript.

“Finally, a key limitation of this approach is that we cannot quantify the proportion of total BK or CaV1.3 channels engaged in hetero-clusters within each compartment. The PLA method provides proximity-based detection, which reflects relative localization rather than absolute channel abundance within individual organelles”.

(5) Temper Overstated Claims

(a) Revise language that suggests the findings introduce a "new paradigm," instead emphasizing how this study extends existing models.

We agree with the reviewers, and we have revised the language to avoid implying a 'new paradigm.' The following is the significance statement.

“This work examines the proximity between BK and CaV1.3 molecules at the level of their mRNAs and newly synthesized proteins to reveal that these channels interact early in their biogenesis. Two cell models were used: a heterologous expression system to investigate the steps of protein trafficking and a pancreatic beta cell line to study the localization of endogenous channel mRNAs. Our findings show that BK and CaV1.3 channels begin assembling intracellularly before reaching the plasma membrane, revealing new aspects of their spatial organization. This intracellular assembly suggests a coordinated process that contributes to functional coupling.”

(b) Moderate conclusions where the supporting data are preliminary or correlative.

We agree with the reviewers, and we have moderated conclusions in instances where the supporting data are preliminary or correlative, ensuring a balanced interpretation. We added the following to the manuscript.

“This study provides novel insights into the organization of BK and CaV1.3 channels in heteroclusters, emphasizing their assembly within the ER, at ER exit sites, and within the Golgi. Our findings suggest that BK and CaV1.3 channels begin assembling intracellularly before reaching the plasma membrane, shaping their spatial organization, and potentially facilitating functional coupling. While this suggests a coordinated process that may contribute to functional coupling, further investigation is needed to determine the extent to which these hetero-clusters persist upon membrane insertion. While our study advances the understanding of BK and CaV1.3 heterocluster assembly, several key questions remain unanswered. What molecular machinery drives this colocalization at the mRNA and protein level? How do disruptions to complex assembly contribute to channelopathies and related diseases? Additionally, a deeper investigation into the role of RNA binding proteins in facilitating transcript association and localized translation is warranted”.

(6) Address Additional Technical and Presentation Issues

(a) Include clearer figure annotations, especially for identifying PLA puncta localization (e.g., membrane vs. intracellular).

We agree with the reviewers, and we have updated the figures to include clearer annotations that distinguish PLA puncta localized at the membrane versus those within intracellular compartments.

(b) Reconsider the scale and arrangement of image panels to better showcase the data.

We agree with the reviewers, and we have adjusted the scale and layout of the image panels to enhance data visualization and readability. Enlarged key regions now provide better clarity of critical features.

(c) Provide precise clone/variant information for BK and CaV1.3 channels used.

We thank the reviewers for their suggestion, and we now provide precise information regarding the BK and CaV1.3 channel constructs used in our experiments, including their Addgene plasmid numbers and relevant variant details. These have been incorporated into the Methods section to ensure reproducibility and transparency. We added the following to the manuscript.

“The CaV1.3 α subunit construct used in our study corresponds to the rat CaV1.3e splice variant containing exons 8a, 11, 31b, and 42a, with a deletion of exon 32. The BK channel construct used in this study corresponds to the VYR splice variant of the mouse BKα subunit (KCNMA1)”.

(d) Correct typographical errors and ensure proper figure/supplementary labeling throughout.

Typographical errors have been corrected, and figure/supplementary labeling has been reviewed for accuracy throughout the manuscript.

(7) Expand the Discussion

(a) Include a brief discussion of findings such as BK surface expression in the absence of CaV1.3.

We thank the reviewers for their suggestion. We expanded the Discussion to include a brief analysis of BK surface expression in the absence of CaV1.3. We included the following in the manuscript.

“BK Surface Expression and Independent Trafficking Pathways

BK surface expression in the absence of CaV1.3 indicates that its trafficking does not strictly rely on CaV1.3-mediated interactions. Since BK channels can be activated by multiple calcium sources, their presence in intracellular compartments suggests that their surface expression is governed by intrinsic trafficking mechanisms rather than direct calcium-dependent regulation. While some BK and CaV1.3 hetero-clusters assemble into signaling complexes intracellularly, other BK channels follow independent trafficking pathways, demonstrating that complex formation is not obligatory for all BK channels. Differences in their transport kinetics further reinforce the idea that their intracellular trafficking is regulated through distinct mechanisms. Studies have shown that BK channels can traffic independently of CaV1.3, relying on alternative calcium sources for activation [13, 41]. Additionally, CaV1.3 exhibits slower synthesis and trafficking kinetics than BK, emphasizing that their intracellular transport may not always be coordinated. These findings suggest that BK and CaV1.3 exhibit both independent and coordinated trafficking behaviors, influencing their spatial organization and functional interactions”.

(b) Clarify why certain colocalization comparisons (e.g., ER vs. ER exit sites) are not directly interpretable.

We thank the reviewer for their suggestion. A clarification has been added to the result section and discussion of the manuscript explaining why colocalization comparisons, such as ER versus ER exit sites, are not directly interpretable. We included the following in the manuscript.

“Result:

ER was not simply due to the extensive spatial coverage of ER labeling, we labeled ER exit sites using Sec16-GFP and probed for hetero-clusters with PLA. This approach enabled us to test whether the hetero-clusters were preferentially localized to ER exit sites, which are specialized trafficking hubs that mediate cargo selection and direct proteins from the ER into the secretory pathway. In contrast to the more expansive ER network, which supports protein synthesis and folding, ER exit sites ensure efficient and selective export of proteins to their target destinations”.

“By quantifying the proportion of BK and CaV1.3 hetero-clusters relative to total channel expression at ER exit sites, we found 28 ± 3% colocalization in tsA-201 cells and 11 ± 2% in INS-1 cells (Figure 3F). While the percentage of colocalization between hetero-clusters and the ER or ER exit sites alone cannot be directly compared to infer trafficking dynamics, these findings reinforce the conclusion that hetero-clusters reside within the ER and suggest that BK and CaV1.3 channels traffic together through the ER and exit in coordination”.

“Colocalization and Trafficking Dynamics

The colocalization of BK and CaV1.3 channels in the ER and at ER exit sites before reaching the Golgi suggests a coordinated trafficking mechanism that facilitates the formation of multi-channel complexes crucial for calcium signaling and membrane excitability [37, 38]. Given the distinct roles of these compartments, colocalization at the ER and ER exit sites may reflect transient proximity rather than stable interactions. Their presence in the Golgi further suggests that posttranslational modifications and additional assembly steps occur before plasma membrane transport, providing further insight into hetero-cluster maturation and sorting events. By examining BK-CaV1.3 hetero-cluster distribution across these trafficking compartments, we ensure that observed colocalization patterns are considered within a broader framework of intracellular transport mechanisms [39]. Previous studies indicate that ER exit sites exhibit variability in cargo retention and sorting efficiency [40], emphasizing the need for careful evaluation of colocalization data. Accounting for these complexities allows for a robust assessment of signaling complexes formation and trafficking pathways”.

Reviewer #1 (Recommendations for the authors):

In addition to the general aspects described in the public review, I list below a few points with the hope that they will help to improve the manuscript:

(1) Page 3: "they bind calcium delimited to the point of entry at calcium channels", better use "sources"

We agree with the reviewer. The phrasing on Page 3 has been updated to use 'sources' instead of 'the point of entry at calcium channels' for clarity.

(2) Page 3 "localized supplies of intracellular calcium", I do not like this term, but maybe this is just silly.

We agree with the reviewer. The term 'localized supplies of intracellular calcium' on Page 3 has been revised to “Localized calcium sources”

(3) Regarding the definitions stated by the authors: How do you distinguish between "ensembles" corresponding to "coordinated collection of BK and Cav channels" and "assembly of BK clusters with Cav clusters"? I believe that hetero-clusters is more adequate. The nomenclature does not respond to any consensus in the protein biology field, and I find that it introduces bias more than it helps. I would stick to heteroclusters nomenclature that has been used previously in the field. Moreover, in some discussion sections, the term "ensemble" is used in ways that border on vague, especially when talking about "functional signaling complexes" or "ensembles forming early." It's still acceptable within context but could benefit from clearer language to distinguish ensemble (structural proximity) from complex (functional consequence).

We agree with the reviewer, and we recognize the importance of precise nomenclature and have adopted hetero-clusters instead of ensembles to align with established conventions in the field. This term specifically refers to the spatial organization of BK and CaV1.3 channels, while functional complexes denote mechanistic interactions. We have revised sections where ensemble was used ambiguously to ensure clear distinction between structure and function.

The definition of "cluster" is clearly stated early but less emphasized in later quantitative analyses (e.g., particle size discussions in Figure 7). Figure 8 is equally confusing, graphs D and E referring to "BK ensembles" and "Cav ensembles", but "ensembles" should refer to combinations of both channels, whereas these seem to be "clusters". In fact, the Figure legend mentions "clusters".

We agree with the reviewer. Terminology has been revised throughout the manuscript to ensure consistency, with 'clusters' used appropriately in quantitative analyses and figure descriptions.

(4) Methods: how are clusters ("ensembles") analysed from the STORM data? What is the logarithm used for? More info about this is required. Equally, more information and discussion about how colocalization is measured and interpreted in superresolution microscopy are required.

We thank the reviewer for their suggestion, and additional details have been incorporated into the Methods section to clarify how clusters ('ensembles') are analyzed from STORM data, including the role of the logarithm in processing. Furthermore, we have expanded the discussion to provide more information on how colocalization is measured and interpreted in super resolution microscopy. We include the following in the manuscript.

“Direct stochastic optical reconstruction microscopy (dSTORM) images of BK and CaV1.3 overexpressed in tsA-201 cells were acquired using an ONI Nanoimager microscope equipped with a 100X oil immersion objective (1.4 NA), an XYZ closed-loop piezo 736 stage, and triple emission channels split at 488, 555, and 640 nm. Samples were imaged at 35°C. For singlemolecule localization microscopy, fixed and stained cells were imaged in GLOX imaging buffer containing 10 mM β-mercaptoethylamine (MEA), 0.56 mg/ml glucose oxidase, 34 μg/ml catalase, and 10% w/v glucose in Tris-HCl buffer. Single-molecule localizations were filtered using NImOS software (v.1.18.3, ONI). Localization maps were exported as TIFF images with a pixel size of 5 nm. Maps were further processed in ImageJ (NIH) by thresholding and binarization to isolate labeled structures. To assess colocalization between the signal from two proteins, binary images were multiplied. Particles smaller than 400 nm2 were excluded from the analysis to reflect the spatial resolution limit of STORM imaging (20 nm) and the average size of BK channels. To examine spatial localization preference, binary images of BK were progressively dilated to 20 nm, 40 nm, 60 nm, 80 nm, 100 nm, and 200 nm to expand their spatial representation. These modified images were then multiplied with the CaV1.3 channel to quantify colocalization and determine BK occupancy at increasing distances from CaV1.3. To ensure consistent comparisons across distance thresholds, data were normalized using the 200 nm measurement as the highest reference value, set to 1”.

(5) Related to Figure 2:

(a) Why use an antibody to label GFP when PH-PLCdelta should be a membrane marker? Where is the GFP in PH-PKC-delta (intracellular, extracellular? Images in Figure 2E are confusing, there is a green intracellular signal.

We thank the reviewer for their feedback. To clarify, GFP is fused to the N-terminus of PH-PLCδ and primarily localizes to the inner plasma membrane via PIP2 binding. Residual intracellular GFP signal may reflect non-membrane-bound fractions or background from anti-GFP immunostaining. We added a paragraph explaining the use of the antibody anti GFP in the Methods section Proximity ligation assay subsection.

(b) The images in Figure 2 do not help to understand how the authors select the PLA puncta located at the plasma membrane. How do the authors do this? A useful solution would be to indicate in Figure 2 an example of the PLA signals that are considered "membrane signals" compared to another example with "intracellular signals". Perhaps this was intended with the current Figure, but it is not clear.

We agree with the reviewer. We have added a sentence to explain how the number of PLA puncta at the plasma membrane was calculated.

“We visualized the plasma membrane with a biological sensor tagged with GFP (PHPLCδ-GFP) and then probed it with an antibody against GFP (Figure 2E). By analyzing the GFP signal, we created a mask that represented the plasma membrane. The mask served to distinguish between the PLA puncta located inside the cell and those at the plasma membrane, allowing us to calculate the number of PLA puncta at the plasma membrane”.

(c) Figure 2C: What is the negative control? Apologies if it is described somewhere, but I seem not to find it in the manuscript.

We thank the reviewer for their suggestion. For the negative control in Figure 2C, BK was probed using the primary antibody without co-staining for CaV1.3 or other proteins, ensuring specificity and ruling out non-specific antibody binding or background fluorescence. A sentence clarifying the negative control for Figure 2C has been added to the Results section, specifying that BK was probed using the primary antibody without costaining for CaV1.3 or other proteins to ensure specificity.

“To confirm specificity, a negative control was performed by probing only for BK using the primary antibody, ensuring that detected signals were not due to non-specific binding or background fluorescence”.

(d) What is the resolution in z of the images shown in Figure 2? This is relevant for the interpretation of signal localization.

The z-resolution of the images shown in Figure 2 was approximately 270–300 nm, based on the Zeiss Airyscan system’s axial resolution capabilities. Imaging was performed with a step size of 300 nm, ensuring adequate sampling for signal localization while maintaining optimal axial resolution.

“In a different experiment, we analyzed the puncta density for each focal plane of the cell (step size of 300 nm) and compared the puncta at the plasma membrane to the rest of the cell”.

(e) % of total puncta in PM vs inside cell are shown for transfected cells, what is this proportion in INS-1 cells?

This quantification was performed for transfected cells; however, we have not conducted the same analysis in INS-1 cells. Future experiments could address this to determine potential differences in puncta distribution between endogenous and overexpressed conditions.

(6) Related to Figure 3:

(a) Figure 3B: is this antibody labelling or GFP fluorescence? Why do they use GFP antibody labelling, if the marker already has its own fluorescence? This should at least be commented on in the manuscript.

We thank the reviewer for their concern. In Figure 3B, GFP was labeled using an antibody rather than relying on its intrinsic fluorescence. This approach was necessary because GFP fluorescence does not withstand the PLA protocol, resulting in significant fading. Antibody labeling provided stronger signal intensity and improved resolution, ensuring optimal signal-to-noise ratio for accurate analysis.

A clarification regarding the use of GFP antibody labeling in Figure 3B has been added to the Methods section, explaining that intrinsic GFP fluorescence does not endure the PLA protocol, necessitating antibody-based detection for improved signal and resolution.We added the following to the manuscript.

“For PLA combined with immunostaining, PLA was followed by a secondary antibody incubation with Alexa Fluor-488 at 2 μg/ml for 1 hour at 21°C. Since GFP fluorescence fades significantly during the PLA protocol, resulting in reduced signal intensity and poor image resolution, GFP was labeled using an antibody rather than relying on its intrinsic fluorescence”.

(b) Why is it relevant to study the ER exit sites? Some explanation should be included in the main text (page 11) for clarification to non-specialized readers. Again, the quantification should be performed on the proportion of clusters/ensembles out of the total number of channels expressed at the ER (or ER exit sites).

We thank the reviewer for their feedback. We have modified this section to include a more detailed explanation of the relevance of ER exit sites to protein trafficking. ER exit sites serve as specialized sorting hubs that regulate the transition of proteins from the ER to the secretory pathway, distinguishing them from the broader ER network, which primarily facilitates protein synthesis and folding. This additional context clarifies why studying ER exit sites provides valuable insights into ensemble trafficking dynamics.

Regarding quantification, our method does not allow for direct measurement of the total number of BK and CaV1.3 channels expressed at the ER or ER exit sites. Instead, we focused on the proportion of hetero-clusters localized within these compartments, which provides insight into trafficking pathways despite the limitation in absolute channel quantification. We included the following in the manuscript in the Results section.

“To determine whether the observed colocalization between BK–CaV1.3 hetero-clusters and the ER was not simply due to the extensive spatial coverage of ER labeling, we labeled ER exit sites using Sec16-GFP and probed for hetero-clusters with PLA. This approach enabled us to test whether the hetero-clusters were preferentially localized to ER exit sites, which are specialized trafficking hubs that mediate cargo selection and direct proteins from the ER into the secretory pathway. In contrast to the more expansive ER network, which supports protein synthesis and folding, ER exit sites ensure efficient and selective export of proteins to their target destinations”.

“By quantifying the proportion of BK and CaV1.3 hetero-clusters relative to total channel expression at ER exit sites, we found 28 ± 3% colocalization in tsA-201 cells and 11 ± 2% in INS-1 cells (Figure 3F). While the percentage of colocalization between hetero-clusters and the ER or ER exit sites alone cannot be directly compared to infer trafficking dynamics, these findings reinforce the conclusion that hetero-clusters reside within the ER and suggest that BK and CaV1.3 channels traffic together through the ER and exit in coordination”.

(7) Related to Figure 4:

A control is included to confirm that the formation of BK-Cav1.3 ensembles is not unspecific. Association with a protein from the Golgi (58K) is tested. Why is this control only done for Golgi? No similar experiment has been performed in the ER. This aspect should be commented on.

We thank the reviewer for their suggestion. We selected the Golgi as a control because it represents the final stage of protein trafficking before proteins reach their functional destinations. If BK and CaV1.3 hetero-cluster formation is specific at the Golgi, this suggests that their interaction is maintained throughout earlier trafficking steps, including within the ER. While we did not perform an equivalent control experiment in the ER, the Golgi serves as an effective checkpoint for evaluating specificity within the broader protein transport pathway. We included the following in the manuscript.

“We selected the Golgi as a control because it represents the final stage of protein trafficking, ensuring that hetero-cluster interactions observed at this point reflect specificity maintained throughout earlier trafficking steps, including within the ER”.

(8) How is colocalization measured, eg, in Figure 6? Are the images shown in Figure 6 representative? This aspect would benefit from a clearer description.

We thank the reviewer for their suggestion. A section clarifying colocalization measurement and the representativeness of Figure 6 images has been added to the Methods under Data Analysis. We included the following in the manuscript.

For PLA and RNAscope experiments, we used custom-made macros written in ImageJ. Processing of PLA data included background subtraction. To assess colocalization, fluorescent signals were converted into binary images, and channels were multiplied to identify spatial overlap.

(9) The text should be revised for typographical errors, for example:

(a) Summary "evidence of" (CHECK THIS ONE)

We agree with the reviewer, and we corrected the typographical errors

(b) Table 1, row 3: "enriches" should be "enrich"

We agree with the reviewer. The term 'enriches' in Table 1, row 3 has been corrected to 'enrich'.

(c) Figure 2B "priximity"

We agree with the reviewer. The typographical errors in Figure 2B has been corrected from 'priximity' to 'proximity'.

(d) Legend of Figure 7 (C) "size of BK and Cav1.3 channels". Does this correspond to individual channels or clusters?

We agree with the reviewer. The legend of Figure 7C has been clarified to indicate that 'size of BK and Cav1.3 channels' refers to clusters rather than individual channels.

(e) Methods: In the RNASCOPE section, "Fig.4-supp1" should be "Fig. 5-supp1"

(f) Page 15, Figure 5B is cited, should be Figure 6B

We agree with the reviewer. The reference in the RNASCOPE section has been updated from 'Fig.4-supp1' to 'Fig. 5-supp1,' and the citation on Page 15 has been corrected from Figure 5B to Figure 6B.

Reviewer #2 (Recommendations for the authors):

(1) The abstract could be more accessible for a wider readership with improved flow.

We thank the reviewer for their suggestion. We modified the summary as follows to provide a more coherent flow for a wider readership.

“Calcium binding to BK channels lowers BK activation threshold, substantiating functional coupling with calcium-permeable channels. This coupling requires close proximity between different channel types, and the formation of BK–CaV1.3 hetero-clusters at nanometer distances exemplifies this unique organization. To investigate the structural basis of this interaction, we tested the hypothesis that BK and CaV1.3 channels assemble before their insertion into the plasma membrane. Our approach incorporated four strategies: (1) detecting interactions between BK and CaV1.3 proteins inside the cell, (2) identifying membrane compartments where intracellular hetero-clusters reside, (3) measuring the proximity of their mRNAs, and (4) assessing protein interactions at the plasma membrane during early translation. These analyses revealed that a subset of BK and CaV1.3 transcripts are spatially close in micro-translational complexes, and their newly synthesized proteins associate within the endoplasmic reticulum (ER) and Golgi. Comparisons with other proteins, transcripts, and randomized localization models support the conclusion that BK and CaV1.3 hetero-clusters form before their insertion at the plasma membrane”.

(2) Figure 2B - spelling of proximity.

We agree with the reviewer. The typographical errors in Figure 2B has been corrected from 'priximity' to 'proximity'.

Reviewer #3 (Recommendations for the authors):

Minor issues to improve the manuscript:

(1) For completeness, the authors should include a few sentences and appropriate references in the Introduction to mention that BK channels are regulated by auxiliary subunits.

We agree with the reviewer. We have revised the Introduction to include a brief discussion of how BK channel function is modulated by auxiliary subunits and provided appropriate references to ensure completeness. These additions highlight the broader regulatory mechanisms governing BK channel activity, complementing the focus of our study. We included the following in the manuscript.

“Additionally, BK channels are modulated by auxiliary subunits, which fine-tune BK channel gating properties to adapt to different physiological conditions. β and γ subunits regulate BK channel kinetics, altering voltage sensitivity and calcium responsiveness [18]. These interactions ensure precise control over channel activity, allowing BK channels to integrate voltage and calcium signals dynamically in various cell types. Here, we focus on the selective assembly of BK channels with CaV1.3 and do not evaluate the contributions of auxiliary subunits to BK channel organization.”

(2) Insert a space between 'homeostasis' and the square bracket at the end of the Introduction's second paragraph.

We agree with the reviewer. A space has been inserted between 'homeostasis' and the square bracket in the second paragraph of the Introduction for clarity.

(3) The images presented in Figures 2-5 should be increased in size (if permitted by the Journal) to allow the reader to clearly see the puncta in the fluorescent images. This would necessitate reconfiguring the figures into perhaps a full A4 page per figure, but I think the quality of the images presented really do deserve to "be seen". For example, Panels A & B could be at the top of Figure 2, with C & D presented below them. However, I'll leave it up to the authors to decide on the most aesthetically pleasing way to show these.

We agree with the reviewer. We have increased the size of Figures 2–8 to enhance the visibility of fluorescent puncta, as suggested. To accommodate this, we reorganized the panel layout for each figure—for example, in Figure 2, Panels A and B are now placed above Panels C and D to support a more intuitive and aesthetically coherent presentation. We believe this revised configuration highlights the image quality and improves readability while conforming to journal layout constraints.

(4) I think that some of the sentences could be "toned down"

(a) eg, in the first paragraph below Figure 2, the authors state "that 46(plus minus)3% of the puncta were localised on intracellular membranes" when, at that stage, no data had been presented to confirm this. I think changing it to "that 46(plus minus)3% of the puncta were localised intracellularly" would be more precise.

(b) Similarly, please consider replacing the wording of "get together at membranes inside the cell" to "co-localise intracellularly".

(c) In the paragraph just before Figure 5, the authors mention that "the abundance of KCNMA1 correlated more with the abundance of CACNA1D than ... with GAPDH." Although this is technically correct, the R2 value was 0.22, which is exceptionally poor. I don't think that the paper is strengthened by sentences such as this, and perhaps the authors might tone this down to reflect this.

(d) The authors clearly demonstrate in Figure 8 that a significant number of BK channels can traffic to the membrane in the absence of Cav1.3. Irrespective of the differences in transcription/trafficking time between the two channel types, the authors should insert a few lines into their discussion to take this finding into account.

We appreciate the reviewer’s feedback regarding the clarity and precision of our phrasing.

Our responses for each point are below.

(a) We have modified the statement in the first paragraph below Figure 2, changing '46 ± 3% of the puncta were localized on intracellular membranes' to '46 ± 3% of the puncta were localized ‘intracellularly’ to ensure accuracy in the absence of explicit data confirming membrane association.

(b) Similarly, we have replaced 'get together at membranes inside the cell' with 'colocalize intracellularly' to maintain clarity and avoid unintended implications.

(c) Regarding the correlation between KCNMA1 and CACNA1D abundance, we recognize that the R² value of 0.22 is relatively low. To reflect this appropriately, we have revised the phrasing to indicate that while a correlation exists, it is modest. We added the following to the manuscript.

“Interestingly, the abundance of KCNMA1 transcripts correlated more with the abundance of CACNA1D transcripts than with the abundance of GAPDH, a standard housekeeping gene, though with a modest R² value.”

(d) To incorporate the findings from Figure 8, we have added discussion acknowledging that a substantial number of BK channels traffic to the membrane independently of CaV1.3. This addition provides context for potential trafficking mechanisms that operate separately from ensemble formation.

(5) For clarity, please insert the word "total" in the paragraph after Figure 3 "..."63{plus minus}3% versus 50%{plus minus}6% of total PLA puncta were localised at the ER". I know this is explicitly stated later in the manuscript, but I think it needs to be clarified earlier.

We agree with the reviewer. The word 'total' has been inserted in the paragraph following Figure 3 to clarify the percentage of PLA puncta localized at the ER earlier in the manuscript

(6) In the discussion, I think an additional (short) paragraph needs to be included to clarify to the reader why the % "colocalization between ensembles and the ER or the ER exit sites can't be compared or used to understand the dynamics of the ensembles". This may permit the authors to remove the last sentence of the paragraph just before the results section, "BK and Cav1.3 ensembles go through the Golgi."

We thank the reviewer for their suggestion. We have added a short paragraph in the discussion to clarify why colocalization percentages between ensembles and the ER or ER exit sites cannot be compared to infer ensemble dynamics. This allowed us to remove the final sentence of the paragraph preceding the results section ('BK and Cav1.3 ensembles go through the Golgi).

(7) In the paragraph after Figure 6, Figure 5B is inadvertently referred to. Please correct this to Figure 6B.

We agree with the reviewer. The reference to Figure 5B in the paragraph after Figure 6 has been corrected to Figure 6B.

(8) In the discussion under "mRNA co-localisation and Protein Trafficking", please insert a relevant reference illustrating that "disruption in mRNA localization... can lead to ion channel mislocalization".

We agree with the reviewer. We have inserted a relevant reference under 'mRNA Colocalization and Protein Trafficking' to illustrate that disruption in mRNA localization can lead to ion channel mislocalization.

(9) The supplementary Figures appear to be incorrectly numbered. Please correct and also ensure that they are correctly referred to in the text.

We agree with the reviewer. The numbering of the supplementary figures has been corrected, and all references to them in the text have been updated accordingly.

(10) The final panels of the currently labelled Figure 5-Supplementary 2 need to have labels A-F included on the image.

We agree with the reviewer. Labels A-F have been added to the final panels of Figure 5-Supplementary 2.

References

(1) Shah, K.R., X. Guan, and J. Yan, Structural and Functional Coupling of Calcium-Activated BK Channels and Calcium-Permeable Channels Within Nanodomain Signaling Complexes. Frontiers in Physiology, 2022. Volume 12 - 2021.

(2) Chen, A.L., et al., Calcium-Activated Big-Conductance (BK) Potassium Channels Traffic through Nuclear Envelopes into Kinocilia in Ray Electrosensory Cells. Cells, 2023. 12(17): p. 2125.

(3) Berkefeld, H., B. Fakler, and U. Schulte, Ca2+-activated K+ channels: from protein complexes to function. Physiol Rev, 2010. 90(4): p. 1437-59.

(4) Loane, D.J., P.A. Lima, and N.V. Marrion, Co-assembly of N-type Ca2+ and BK channels underlies functional coupling in rat brain. J Cell Sci, 2007. 120(Pt 6): p. 98595.

(5) Boncompain, G. and F. Perez, The many routes of Golgi-dependent trafficking. Histochemistry and Cell Biology, 2013. 140(3): p. 251-260.

(6) Kurokawa, K. and A. Nakano, The ER exit sites are specialized ER zones for the transport of cargo proteins from the ER to the Golgi apparatus. The Journal of Biochemistry, 2019. 165(2): p. 109-114.

(7) Chen, G., et al., BK channel modulation by positively charged peptides and auxiliary γ subunits mediated by the Ca2+-bowl site. Journal of General Physiology, 2023. 155(6).

Changes from the earlier version to the revised version

Abstract

Earlier version: Calcium binding to BK channels lowers their activation threshold. Hence, BK channels are functionally coupled to calcium-permeable channels. This functional coupling requires proximity of two different types of channels. Formation of an ensemble at nanometer distances between BK and CaV1.3 channels exemplifies this unique organization. We investigated the mechanism underlying their structural organization, testing the hypothesis that the assembly between BK and CaV1.3 channels is formed before their insertion at the plasma membrane. Our design used four approaches: 1) to detect the interaction between BK and CaV1.3 proteins inside the cell, 2) to identify the membranes where intracellular ensembles reside, 3) to measure the proximity between their mRNAs, and 4) to look for protein interaction at the plasma membrane soon after the start of their translation. These measures showed that a percentage of transcripts for BK and CaV1.3 channels are close and that the newly synthesized proteins interact through the endoplasmic reticulum and Golgi. Comparisons were made with other proteins and transcripts, as well as with predictions of randomized localizations of BK channels. We found evidence of assembly between BK and CaV1.3 ensembles before their insertion at the plasma membrane.

Revised version: Calcium binding to BK channels lowers BK activation threshold, substantiating functional coupling with calcium-permeable channels. This coupling requires close proximity between different channel types, and the formation of BK–CaV1.3 hetero-clusters at nanometer distances exemplifies this unique organization. To investigate the structural basis of this interaction, we tested the hypothesis that BK and CaV1.3 channels assemble before their insertion into the plasma membrane. Our approach incorporated four strategies: (1) detecting interactions between BK and CaV1.3 proteins inside the cell, (2) identifying membrane compartments where intracellular hetero-clusters reside, (3) measuring the proximity of their mRNAs, and (4) assessing protein interactions at the plasma membrane during early translation. These analyses revealed that a subset of BK and CaV1.3 transcripts are spatially close in micro-translational complexes, and their newly synthesized proteins associate within the endoplasmic reticulum (ER) and Golgi. Comparisons with other proteins, transcripts, and randomized localization models support the conclusion that BK and CaV1.3 hetero-clusters form before their insertion at the plasma membrane.

Introduction

This work concerns the organization of functionally coupled voltage- and calcium-dependent potassium (BK) channels and voltage-gated calcium (CaV) channels.

Earlier version: BK channels (a.k.a. KCa1.1, Maxi-K, Slo1, KCNMA1) are named for their “big potassium (K+)” conductance, since channel opening leads to unusually large, outward potassium currents and membrane repolarization. Their amino acid sequence is conserved throughout the animal kingdom from worms to mammals. BK channels are expressed in a wide variety of cell types, predominantly excitable cells but also non-excitable salivary, bone and kidney cells, making this channel responsible for a large range of physiological processes [1-9]. Phenotypes of pathological BK mutations in human patients are most prominent in the brain and muscle, and often manifest with seizures, movement disorders, developmental delay, and intellectual disability [10, 11]. BK is also implicated in other organ system functions, including but not limited to heart pace making, reproduction, and pancreatic glucose homeostasis[11, 12].

Revised version: BK channels (a.k.a. KCa1.1, Maxi-K, Slo1, KCNMA1) are named for their “big potassium (K+)” conductance, since channel opening leads to unusually large, outward potassium currents and membrane repolarization. Its amino acid sequence is conserved throughout the animal kingdom from worms to mammals. BK channels are expressed in a wide variety of cell types, predominantly excitable cells but also non-excitable salivary, bone and kidney cells, making this channel responsible for a large range of physiological processes [1-9]. Phenotypes of pathological BK mutations in human patients are most prominent in the brain and muscle and often manifest as seizures, movement disorders, developmental delay, and intellectual disability [10, 11]. BK is also implicated in other organ system functions, including but not limited to heart pace making, reproduction, and pancreatic glucose homeostasis [11, 12].

Earlier version: Since BK channels affect function in numerous physiological systems, their channel opening is meticulously regulated at the cellular level. BK is a voltage-gated channel, activated by membrane depolarization. Interestingly, BK has an additional gating mechanism that differentiates it from typical voltage-gated potassium channels; BK opening is gated by both voltage and intracellular calcium binding to a cytosolic regulatory domain. At the cytoplasmic resting free calcium concentration (∼0.1 µM), BK channels remain closed [13]. However, the probability of opening increases both when the membrane potential depolarizes and when the local free calcium rises. Yet, increases of calcium are tightly confined by endogenous proteins that bind calcium with high affinity and extrusion mechanisms that take calcium inside organelles or outside the cell. Hence, BK activation relies on strategies to overcome these barriers. They bind calcium delimited to the point of entry at calcium channels.

Revised version: Since BK channels affect function in numerous physiological systems, channel opening is meticulously regulated at the cellular level. BK is a voltage-gated channel activated by membrane depolarization. Interestingly, BK has an additional gating mechanism that differentiates it from typical voltage-gated potassium channels; BK opening is gated by both voltage and intracellular calcium binding to a cytosolic regulatory domain. At the cytoplasmic resting free calcium concentration (∼0.1 µM), BK channels remain closed [13]. However, the probability of channel opening increases when both the membrane potential depolarizes and when the local free calcium rises. Yet, increases of calcium are tightly limited by endogenous proteins that bind calcium with high affinity as well as by extrusion mechanisms that take calcium inside organelles or outside the cell. Hence, the activation of BK channels relies on strategies to overcome these regulatory barriers. One such strategy is to localize near sources of calcium.

Earlier version: In excitable cells, BK forms nanodomains with calcium channels that provide exclusive, localized supplies of intracellular calcium. Several subtypes of calcium channels form functional signaling complexes with BK, shifting BK activation voltages to more negative potentials [14-16]. One of these channels, CaV1.3 (a.k.a. CACNA1D), is unique in its electrophysiological profile as an L-type calcium channel, activating with fast kinetics at voltages as negative as -55 mV [17].

Revised version: In excitable cells, BK forms nanodomains with calcium channels that provide exclusive, localized calcium sources. Several subtypes of calcium channels form functional signaling complexes with BK, shifting BK activation voltages to more negative potentials [14-16]. One of these channels, CaV1.3 (a.k.a. CACNA1D), is unique in its electrophysiological profile as an L-type calcium channel, activating with fast kinetics at voltages as negative as −55 mV [17]. Additionally, BK channels are modulated by auxiliary subunits, which fine-tune BK channel gating properties to adapt to different physiological conditions. β and γ subunits regulate BK channel kinetics, altering voltage sensitivity and calcium responsiveness [18]. These interactions ensure precise control over channel activity, allowing BK channels to integrate voltage and calcium signals dynamically in various cell types. Here, we focus on the selective assembly of BK channels with CaV1.3 and do not evaluate the contributions of auxiliary subunits to BK channel organization. CaV1.3 is expressed in many of the same cell types as BK and is often functionally coupled with BK channels, enabling BK channels to activate at more negative voltages. Notably, super-resolution microscopy shows that CaV1.3 organizes spatially into nanodomains with BK in the plasma membrane [15]. However, the mechanisms behind the assembly of BK and CaV1.3 hetero-clusters are unknown.

Removed in the revised version: CaV1.3 is expressed in many of the same cell types as BK and is often functionally coupled with BK channels, enabling them to activate at more negative voltages. Notably, super-resolution microscopy shows that CaV1.3 organizes spatially into nanodomains near BK in the plasma membrane [15]. However, the mechanisms behind the assembly of BK and CaV1.3 complexes are unknown.

Earlier version: Several mechanisms for bringing ion channels together have been suggested (table 1). One proposes that it precedes protein insertion into the plasma membrane [18-20]. This mechanism has been observed in studies of hetero-multimers and even channels permeating different ions. The formation of these groups of proteins would start from their synthesis; possibly, even their mRNA transcripts are colocalized prior to translation. We explored this mechanism in relation to BK and CaV1.3 functional coupling. Here, we started by asking when BK and CaV1.3 channels cluster in the cell. We looked for BK and CaV1.3 ensembles at intracellular membranes of the ER, at ER exit sites, and at the Golgi. We also investigated the proximity between mRNA transcripts of BK and CaV1.3.

Revised version: Several mechanisms for bringing ion channels together have been suggested (Table 1). One proposes that proteins assembly precedes protein insertion into the plasma membrane [19, 25, 26]. This mechanism has been observed in studies of hetero-multimers and even channels permeating different ions. The formation of these groups of proteins would start from their synthesis. There is even the possibility of mRNA transcripts colocalizing prior to translation. We explored this mechanism in relation to BK and CaV1.3 functional coupling. Here, we started by investigating when and where BK and CaV1.3 channels cluster in the cell. We looked for BK and CaV1.3 hetero-clusters at intracellular membranes of the ER, at ER exit sites, and at the Golgi. We also investigated the proximity between mRNA transcripts of BK and CaV1.3.

Definitions used in this manuscript

Earlier version: To guide the reader and prevent confusion on the terms used, we introduce the following nomenclature, which is also illustrated in Figure 1. Molecular complex: an array of several polypeptides with a defined function. In our case, a BK channel is a molecular complex of four alpha subunits whose function is to permeate ions, also referred as a homo-tetramer. Cluster: we define a cluster as the accumulation of proteins. In our work, cluster refers to an accumulation of BK channels or an accumulation of calcium channels. In a previous manuscript, we also referred to these as homo-clusters. Ensemble: we define ensemble as the collection of different types of proteins that facilitate functional coupling and compartmentalization of a signaling complex. These proteins can also be in clusters; in such cases, an ensemble is the aggregation of clusters. In the present work, ensemble refers to a coordinated collection of BK and CaV channels. It is also used to refer to the assembly of BK clusters with CaV clusters. In a previous manuscript, we referred to this as hetero-clusters.

Revised version: To guide the reader and prevent confusion on the terms used, we introduce the following nomenclature, which is also illustrated in Figure 1. Molecular complex: an array of several polypeptides with a defined function. In our case, a BK channel is a molecular complex of four alpha subunits whose function is to permeate ions. Homo-cluster: we define a homo-cluster as the accumulation of proteins. In our work, a homo-cluster refers to the accumulation of BK channels or the accumulation of calcium channels.

Earlier version: Representation of molecular complex, cluster, and ensemble.

Revised version: Representation of molecular complex, homo-cluster, and hetero-cluster.

Added in the revised version: Hetero-cluster: we define a hetero-cluster as the collection of different types of proteins that facilitate functional coupling and compartmentalization of a signaling complex. In the present work, hetero-cluster refers to a coordinated collection of BK and CaV channels. It is also used to refer to the assembly of homo-clusters of BK with homo-clusters of CaV channels.

Materials and Methods

Cell culture

Earlier version: We used tsA-201 cells to co-express BK and CaV1.3 channels heterologously. Cells were grown in DMEM (Gibco) supplemented with 10% fetal bovine serum and 0.2% penicillin/streptomycin. We used rat insulinoma (INS-1) cells to study endogenous levels of transcripts and proteins of the channels. INS-1 cells were cultured in RPMI high glutamate medium (Gibco) with 10% fetal bovine serum, 0.2% penicillin/streptomycin, 10 mM HEPES (Gibco), 1 mM sodium pyruvate (Gibco), and 50 μM 2-mercaptoethanol. Both cell types were passaged twice a week and incubated in 5% CO2 at 37°C.

Revised version: We used tsA-201 cells to co-express BK and CaV1.3 channels heterologously. Cells were grown in DMEM (Gibco) supplemented with 10% fetal bovine serum and 0.2% penicillin/streptomycin. We used rat insulinoma (INS-1) cells to study endogenous levels of transcripts and proteins of channels. INS-1 cells were cultured in RPMI high glutamate medium (Gibco) with 10% fetal bovine serum, 0.2% penicillin/streptomycin, 10 mM HEPES (Gibco), 1 mM sodium pyruvate (Gibco), and 50 μM 2-mercaptoethanol. Both cell types were passaged twice a week and incubated in 5% CO2 at 37° C.

Plasmids and transfection

Earlier version: Cells were transfected with 0.1– 0.4 μg DNA per plasmid, plated for 24 hours on poly-D-Lysine coated coverslips, and used for experiments. Lipofectamine 3000 (Invitrogen, RRID: L30000) was used for the transfection. DNA clones of CaV1.3, BK channels, PH-PLCδ GFP, ER moxGFP, pmGFP-Sec16S and Golgi-mGFP were obtained from Addgene (RRID:SCR_002037). Auxiliary subunits for CaV1.3 channels, CaVβ3 and CaVα2δ1 (from Diane Lipscombe, Brown University, RI), were transfected as well. No BK channel auxiliary subunits were transfected.

Revised version: Cells were transfected with 0.1– 0.4 μg DNA per plasmid and plated for 24 hours on poly-D-Lysine coated coverslips. Lipofectamine 3000 (Invitrogen, RRID: L30000) was used for the transfection. DNA clones of CaV1.3, BK channels, PH-PLCδ GFP, ER moxGFP, pmGFP-Sec16S and Golgi-mGFP were obtained from Addgene (RRID:SCR_002037). The CaV1.3 α subunit construct used in our study corresponds to the rat CaV1.3e splice variant containing exons 8a, 11, 31b, and 42a, with a deletion of exon 32. The BK channel construct corresponds to the VYR splice variant of the mouse BKα subunit (KCNMA1). Auxiliary subunits for CaV1.3 channels, CaVβ3 and CaVα2δ1 (from Diane Lipscombe, Brown University, RI), were transfected as well. No BK channel auxiliary subunits were transfected.

Antibodies

Earlier version: CaV1.3 channels were immuno-detected with a rabbit primary antibody recognizing residues 809 to 825 located at the intracellular II-III loop of the channel (DNKVTIDDYQEEAEDKD), kindly provided by Drs. William Catterall and Ruth Westenbroek [25]. BK channels were detected using the anti-Slo1 mouse monoclonal antibody clone L6/60. The goat polyclonal GFP antibody was against the recombinant full-length protein corresponding to Aequorea victoria GFP. Anti-58K Golgi protein antibody was used to mark the Golgi. Specificity of antibodies was tested in un-transfected tsA-201 cells (Figure 1-supplementary 1). The secondary antibodies tagged with Alexa dyes were Donkey anti-mouse Alexa-647, Donkey anti-rabbit Alexa-555, Donkey anti-goat Alexa-488, all from Molecular Probes.

Revised version: CaV1.3 channels were immuno-detected with a rabbit primary antibody recognizing residues 809 to 825 located at the intracellular II-III loop of the channel (DNKVTIDDYQEEAEDKD), kindly provided by Drs. William Catterall and Ruth Westenbroek [27]. BK channels were detected using the anti-Slo1 mouse monoclonal antibody clone L6/60. The goat polyclonal GFP antibody was against the recombinant full-length protein corresponding to Aequorea victoria GFP. Anti-58K Golgi protein antibody was used to mark the Golgi. Specificity of antibodies was tested in un-transfected tsA-201 cells (Figure 2-Supplementary 1). The secondary antibodies tagged with Alexa dyes were Donkey anti-mouse Alexa-647, Donkey anti-rabbit Alexa-555, Donkey anti-goat Alexa-488 (Molecular Probes).

Added in the revised version: BK and CaV1.3 hetero-clusters are found inside the cell.

Added in the revised version: A. Diagram of the hypothesis: hetero-clusters of BK (magenta) and CaV1.3 (cyan) are on intracellular membranes and on the plasma membrane. B. Illustration of the technique to detect BK and CaV1.3 hetero-clusters. Proximity ligation assay is used to detect the hetero-clusters. C. Confocal images of fluorescent puncta from PLA experiments in tsA-201 cells. Left: Cells were transfected and probed for BK and CaV1.3 channels. Right: negative control. Cells were transfected and probed only for BK channels. Enlargement of a selected region is shown in the inset. D. Scatter dot plot comparing puncta density of BK and CaV1.3 hetero-clusters to the negative control. Data points are from n = 12 cells for BK and CaV1.3 hetero-clusters and from n = 14 cells for negative control. p-values are shown at the top of the graphs. E. Confocal images of fluorescent PLA puncta at different focal planes co-labeled against GFP at the plasma membrane. Cells were transfected with BK, CaV1.3, and PH-PLCδ-GFP and probed for BK channels, CaV1.3 channels, and GFP. PLA puncta are shown in magenta, and the plasma membrane is shown in green. Enlargements of the representative regions of PM and intercellular hetero-clusters are shown in the insets. F. Scatter dot plot comparing BK and CaV1.3 hetero-cluster abundance at PM and inside the cell. Data points are from n = 12 cells. Scale bar is 10 μm and 1 μm in the insets.

Added in the revised version: Validation of antibodies against BK, CaV1.3 and GFP.

Added in the revised version: Representative confocal images of tsA-201 cells immuno-tested for BK channels (left), CaV1.3 channels (middle), and GFP proteins (right). Cells were not transfected. Nuclei were stained with DAPI and pseudo colored in gray. Scale bar is 10 μm for all images.

Immunostaining

Earlier version: Cells were fixed with freshly prepared 4% paraformaldehyde for 10 min. After washing, aldehydes were reduced with 0.1% NaBH4 for 5 minutes at 21°C and then washed again. Nonspecific binding was blocked with 3% bovine serum albumin (Thermo Scientific). Cells were permeabilized with 0.25% v/v Triton X-100 in PBS for 1 hour. Primary antibodies were used at 10 μg/ml in blocking solution and incubated overnight at 4°C. After washing, secondary antibodies at 2 μg/ml were incubated for 1 hour at 21°C. For single-molecule localization microscopy, fixed and stained cells were imaged in GLOX imaging buffer containing 10 mM β-mercaptoethylamine (MEA), 0.56 mg/ml glucose oxidase, 34 μg/ml catalase, and 10% w/v glucose in Tris-HCl buffer. Washing steps indicated in all the methods in this paper include 3 cycles of rinses and rocking for 5 minutes with PBS. Cells were imaged using an inverted AiryScan microscope or an ONI Nanoimager with super-resolution capabilities.

Revised version: Cells were fixed with freshly prepared 4% paraformaldehyde for 10 minutes. After washing, aldehydes were reduced with 0.1% NaBH4 for 5 minutes and then washed again. Nonspecific binding was blocked with 3% bovine serum albumin (Thermo Scientific). Cells were permeabilized with 0.25% v/v Triton X-100 in PBS for 1 hour. Primary antibodies were used at 10 μg/ml in blocking solution and incubated overnight at 4° C. After washing, secondary antibodies at 2 μg/ml were incubated for 1 hour at 21°C. Washing steps indicated in all methods include 3 cycles of rinsing and rocking for 5 minutes with PBS at 21°C. Cells were imaged using an inverted AiryScan microscope or an ONI Nanoimager with super-resolution capabilities, in total internal reflection fluorescence (TIRF) mode, and with a Z-resolution of 50 nm.

Proximity ligation assay (PLA)

Earlier version: Cells were fixed with freshly prepared 4% paraformaldehyde for 10 minutes. After washing, aldehydes were reduced with 50 mM glycine for 15 minutes at 21°C. After another round of washes, PLA was performed according to manufacturer instructions (Duolink® In Situ Red Starter Kit). Cells were blocked and permeabilized with Duolink blocking solution. Primary antibodies were used at 10 μg/ml in Duolink antibody diluent and incubated overnight at 4°C. The Duolink® In Situ PLA® probe anti-rabbit PLUS and anti-mouse MINUS were used as secondary antibodies, followed by ligation and amplification. For PLA combined with immunostaining, PLA was followed by a secondary antibody incubation with Alexa Fluor-488 (Molecular Probes) at 2 μg/ml for 1 hour at 21°C. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. Cells were imaged using an inverted Zeiss AiryScan microscope.

Revised version: Cells were fixed with freshly prepared 4% paraformaldehyde for 10 minutes. After washing, aldehydes were reduced with 50 mM glycine for 15 minutes. After another round of washes, PLA was performed according to manufacturer instructions (Duolink® In Situ Red Starter Kit). Cells were blocked and permeabilized with Duolink blocking solution. Primary antibodies were used at 10 μg/ml in Duolink antibody diluent and incubated overnight at 4° C. The Duolink® In Situ PLA® probe anti-rabbit PLUS and anti-mouse MINUS were used as secondary antibodies, followed by ligation and amplification. For PLA combined with immunostaining, PLA was followed by a secondary antibody incubation with Alexa Fluor-488 at 2 μg/ml for 1 hour at 21°C. Since GFP fluorescence fades significantly during the PLA protocol, resulting in reduced signal intensity and poor image resolution, GFP was labeled using an antibody rather than relying on its intrinsic fluorescence. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. Cells were imaged using an inverted Zeiss AiryScan microscope.

Earlier version: Single molecule In Situ Hybridization (RNASCOPE)

Revised version: Single-molecule fluorescence in situ hybridization (RNAscope™)

Earlier version: Manual RNAscope assay was performed using RNAscope™ Multiplex Fluorescent V2 Assay according to the manufacturer’s protocol. The RNAScope assay consists of target probes and a signal amplification system composed of a preamplifier, amplifier, and label probe. A schematic RNAscope assay procedure is shown in Figure 4-Supplementary 1. Briefly, Cells were fixed with 4% paraformaldehyde for 30 minutes, washed, dehydrated, and then rehydrated with ethanol, and permeabilized with 0.1% Tween-20 in PBS. Next, cells were quenched with H2O2 and treated with Protease III. Probes were hybridized for 2 hours at 40°C followed by RNAscope amplification and then fluorescent detection. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. We used the following RNAscope probes: RNAscope 3-plex Positive Control Probes, RNAscope 3-plex negative control probes, RNAscope™ Probe-Rn-Ryr, RNAscope™ Probe-Rn-Scn9a, RNAscope Probe-Rn-Kcnma1-C3, RNAscope Probe-Rn-Cacna1d-C2 and RNAscope Probe-Rn-Gapdh. Cells were imaged on the inverted AiryScan microscope.

Revised version: Manual RNAscope assay was performed using RNAscope Multiplex Fluorescent V2 Assay according to the manufacturer’s protocol. The RNAscope assay consists of target probes and a signal amplification system composed of a preamplifier, amplifier, and label probe. A schematic RNAscope assay procedure is shown in Figure 5-Supplementary 1. Briefly, cells were fixed with 4% paraformaldehyde for 30 minutes, washed, dehydrated, and then rehydrated with ethanol, and permeabilized with 0.1% Tween-20 in PBS. Next, cells were quenched with H2O2 and treated with Protease III. Probes were hybridized for 2 hours at 40° C followed by RNAscope amplification and then fluorescence detection. Coverslips were mounted using ProLong™ Gold Antifade Mountant with DAPI. We used the following RNAscope probes: RNAscope 3-plex Positive Control Probes, RNAscope 3-plex negative control probes, RNAscope™ Probe-Rn-Ryr, RNAscope™ Probe-Rn-Scn9a, RNAscope Probe-Rn-Kcnma1-C3, RNAscope Probe-Rn-Cacna1d-C2, and RNAscope ProbeRn-Gapdh. Cells were imaged on the inverted AiryScan microscope. For PLA and RNAscope experiments, we used custom-made macros written in ImageJ. Processing of PLA data included background subtraction. To assess colocalization, fluorescent signals were converted into binary images, and channels were multiplied to identify spatial overlap. Specificity of RNAscope probes was tested in un-transfected tsA-201 cells (Figure 5-Supplementary 2). For RNAscope combined with immunostaining, RNAscope was followed by blocking in PBS supplemented with 0.01 % Tween-20 and 3% BSA for 1 h at 21°C. Samples were then probed for BK protein using primary antibody overnight at 4°C followed by secondary antibody incubation with Alexa Fluor-488 at 2 μg/ml for 1 hour at 21°C. Coverslips were mounted using ProLong™ Gold Antifade Mountant. Cells were imaged using an inverted Zeiss AiryScan microscope.

High resolution imaging

Cells were imaged using an inverted AiryScan microscope (Zeiss LSM 880) run by ZEN black v2.3 software and equipped with a plan apochromat 63X oil immersion objective with 1.4 NA. Fluorescent dyes were excited with a 405 nm diode, 458–514 nm argon, 561 nm, or 633 nm laser. Emission light was detected using an Airyscan 32 GaAsP detector and appropriate emission filter sets. The point spread functions were calculated using ZEN black software and 0.1 μm fluorescent microspheres. The temperature inside the microscope housing was 22° C. Images were analyzed using ImageJ (NIH).

Super-resolution imaging

Earlier version: Direct stochastic optical reconstruction microscopy (dSTORM) images of BK and CaV1.3 overexpressed in tSA-201 cells were obtained using ONI Nanoimager microscope equipped with a 100X oil immersion objective with a 1.4 NA, an XYZ closed-loop piezo 736 stage, and triple emission channels split at 488, 555 and 640 nm. Samples were imaged in GLOX imaging buffer at 35°C. The acquired single-molecule data were filtered using NImOS software (v.1.18.3, ONI) and further processed in ImageJ (NIH).

Revised version: Direct stochastic optical reconstruction microscopy (dSTORM) images of BK and CaV1.3 overexpressed in tsA-201 cells were acquired using an ONI Nanoimager microscope equipped with a 100X oil immersion objective (1.4 NA), an XYZ closed-loop piezo 736 stage, and triple emission channels split at 488, 555, and 640 nm. Samples were imaged at 35° C. For single-molecule localization microscopy, fixed and stained cells were imaged in GLOX imaging buffer containing 10 mM β-mercaptoethylamine (MEA), 0.56 mg/ml glucose oxidase, 34 μg/ml catalase, and 10% w/v glucose in Tris-HCl buffer. Single-molecule localizations were filtered using NImOS software (v.1.18.3, ONI). Localization maps were exported as TIFF images with a pixel size of 5 nm. Maps were further processed in ImageJ (NIH) by thresholding and binarization to isolate labeled structures. To assess colocalization between the signal from two proteins, binary images were multiplied. Particles smaller than 400 nm2 were excluded from the analysis to reflect the spatial resolution limit of STORM imaging (20 nm) and the average size of BK channels. To examine spatial localization preference, binary images of BK were progressively dilated to 20 nm, 40 nm, 60 nm, 80 nm, 100 nm, and 200 nm to expand their spatial representation. These modified images were then multiplied with the CaV1.3 channel to quantify colocalization and determine BK occupancy at increasing distances from CaV1.3. To ensure consistent comparisons across distance thresholds, data were normalized using the 200 nm measurement as the highest reference value, set to 1.

Image scrambling

Earlier version: Images were binarized as .tif images, and their respective cell perimeter coordinates were exported as .csv files by ImageJ (NIH). Processed binary images were then analyzed by our SpotScrambler (https://github.com/jehuang2/SpotScramble) Python program. SpotScrambler first extracts the areas of fluorescent particles in the binary image. SpotScrambler then redraws the fluorescent particles as circles at randomized coordinates within cell perimeter boundaries. SpotScrambler accurately preserves particle number and particle sizes, averaging less than 1% difference in total area of fluorescent particles between pre-SpotScrambler and post-SpotScrambler images. To ensure reliable randomization for each experiment, results were averaged between 3 trials of SpotScrambler.

Revised version: Images were binarized as TIFF images, and their respective cell perimeter coordinates were exported as CSV files by ImageJ. Processed binary images were then analyzed by our SpotScrambler (https://github.com/jehuang2/SpotScrambler) Python program. SpotScrambler first extracts the areas of fluorescent particles in the binary image. SpotScrambler then redraws the fluorescent particles as circles at randomized coordinates within cell perimeter boundaries. SpotScrambler accurately preserves particle number and particle sizes, averaging less than 1% difference in total area of fluorescent particles between pre-SpotScrambler and post-SpotScrambler images. To ensure reliable randomization for each experiment, results were averaged between 3 trials of SpotScrambler.

Data analysis

Earlier version: Excel (Microsoft), and Prism (GraphPad) were used to analyze data. ImageJ was used to process images. For PLA and RNAScope experiments we used custom-made macros written in ImageJ. Localization maps acquired from super-resolution experiments were reconstructed using NImOS software. Processing consisted of thresholding and binarization of images to isolate labeled structures. Particles were analyzed to calculate both the area and number of particles and the colocalization from these images. Processing of data from PLA experiments included background subtraction. One-way ANOVA and non-parametric statistical test (Mann-Whitney Wilcoxon) were used to test for statistical significance. p-values <0.05 were deemed statistically significant. The number of cells used for each experiment is detailed in each figure legend.

Revised version: Excel (Microsoft), and Prism (GraphPad) were used to analyze data. ImageJ was used to process images. One-way ANOVA and non-parametric statistical test (Mann-Whitney Wilcoxon) were used to test for statistical significance. p-values <0.05 were deemed statistically significant. The number of cells used for each experiment is detailed in each figure legend.

Results

Earlier version: BK and CaV1.3 ensembles are found inside the cell

Revised version: BK and CaV1.3 hetero-clusters are found inside the cell

Earlier version: BK and CaV ensembles have been observed at the plasma membrane [14-16]. However, a clear understanding of when or how BK and CaV complexes assemble is missing. A simple mechanism supposes that these ensembles organize only at the plasma membrane. Here, we tested an alternative hypothesis that BK and CaV assemble inside the cell (Figure 2A). To test this idea, we used proximity ligation assay (PLA) and antibodies against BK and CaV1.3 channels. Only when these antibodies are closer than 40 nm, can ligation and amplification occur, resulting in the formation of fluorescent puncta (Figure 2B), here referred to as PLA puncta. Hence, the PLA puncta in Figure 2C represent BK and CaV1.3 ensembles. We first analyzed Z-projections, in which all the puncta in the cell volume are added, and found a density of 5.8 ± 1.0 / 10 μm2. In a different experiment, we analyzed the puncta density for each focal plane of the cell and compared the puncta at the plasma membrane to the rest of the cell. We visualized the plasma membrane with a biological sensor tagged with GFP (PH-PLCδ-GFP) and probed against GFP using an antibody (Figure 2E). We identified clear puncta at the plasma membrane, but to our surprise, also puncta inside the cell. We found that 46 ± 3 % of the puncta were localized on intracellular membranes, whereas 54 ± 3 % were at the plasma membrane. This finding is consistent with our supposition that BK and CaV1.3 channels get together at membranes inside the cell.

Revised version: BK and CaV hetero-clusters have been observed at the plasma membrane [14-16]. However, a clear understanding of when or how BK and CaV hetero-clusters assemble is missing. A simple mechanism proposes that these hetero-clusters organize only at the plasma membrane. Here, we tested the alternative hypothesis of BK and CaV assembling inside the cell (Figure 2A). To test this idea, we used proximity ligation assay (PLA) and antibodies against BK and CaV1.3 channels. When these antibodies are within 40 nm of each other, PLA ligation and amplification can occur, resulting in the formation of fluorescent puncta (Figure 2B), here referred to as PLA puncta. Hence, the PLA puncta in Figure 2C represent BK and CaV1.3 hetero-clusters. To confirm specificity, a negative control was performed by probing only for BK using the primary antibody, ensuring that detected signals were not due to non-specific binding or background fluorescence. We first analyzed Z-projections, in which all the puncta in the cell volume are added, and found a density of 5.8 ± 1.0 / 10 μm2. In a different experiment, we analyzed the puncta density for each focal plane of the cell (step size of 300 nm) and compared the puncta at the plasma membrane to the rest of the cell. We visualized the plasma membrane with a biological sensor tagged with GFP (PH-PLCδ-GFP) and then probed it with an antibody against GFP (Figure 2E). By analyzing the GFP signal, we created a mask that represented the plasma membrane. The mask served to distinguish between the PLA puncta located inside the cell and those at the plasma membrane, allowing us to calculate the number of PLA puncta at the plasma membrane. To our surprise, we found a significant number of puncta localized inside the cell. 46 ± 3 % of the puncta were localized intracellularly, whereas 54 ± 3 % were at the plasma membrane. This finding is consistent with our supposition that BK and CaV1.3 channels colocalize in the cell.

Removed in the revised version: BK and CaV1.3 ensembles are found inside the cell.

Removed in the revised version: A. Diagram of the hypothesis: Ensembles of BK (magenta) and CaV1.3 (cyan) are on intracellular membranes and on the plasma membrane. B. Illustration of the technique to detect BK and CaV1.3 ensembles. Proximity ligation assay is used to detect the ensembles. C. Confocal images of fluorescent puncta from PLA experiments in tsA-201 cells. Left: Cells were transfected and probed for BK and CaV1.3 channels. Right: negative control. Cells were transfected and probed only for BK channels. Enlargement of a selected region is shown in the inset. D. Scatter dot plot comparing puncta density of ensembles of BK and CaV1.3 to the negative control. Data points are from n = 12 cells from BK and CaV1.3 clusters and from n = 14 cells from negative control. p-values are shown at the top of the graphs. E. Confocal images of fluorescent PLA puncta at different focal planes co-labeled against GFP at the plasma membrane. Cells were transfected with BK, CaV1.3, and PH-PLCδ-GFP and probed for BK and CaV1.3 channels and GFP. The plane of the PM was set as zero in the Z axis. PLA puncta are shown in magenta, and the plasma membrane in green. Distance from the PM plane is labeled with the Z values on top of each image. Enlargements of the same region in Z are shown in the insets. F. Scatter dot plot comparing BK and CaV1.3 cluster abundance at PM and inside the cell. Data points are from n = 12 cells. Scale bar is 10 μm and 1 μm in the insets.

Removed in the revised version: Validation of antibodies against BK, CaV1.3 and GFP.

Removed in the revised version: Representative confocal images of tsA-201 cells immuno-tested for BK channels (left), CaV1.3 channels (middle), and GFP proteins (right). Cells were not transfected. Nuclei were stained with DAPI and pseudo colored in gray. Scale bar is 10 μm for all images.

Earlier version: BK and CaV1.3 ensembles localize at ER and ER exit sites

Revised version: BK and CaV1.3 hetero-clusters localize at ER and ER exit sites

Earlier version: We next investigated the identity of the intracellular membranes where these PLA puncta were found. A large component of the intracellular membranes is the endoplasmic reticulum (ER), where channels are inserted after translation. To determine whether BK and CaV associate in the ER (Figure 3A), we combined PLA with immunodetection. As before, PLA probed for ensemble formation between BK and CaV channels, whereas a KDEL-moxGFP label identified the ER (Figure 3B). To avoid disruption of organelle architecture, we used the monomeric mox version of KDEL-GFP, which is optimized to reduce oligomerization in the ER environment [26]. Neither overexpression of KDEL-moxGFP nor fixation altered the ER structure as the tubule width remained around 150 nm for either condition (Figure 3C), a value in agreement with the literature [27, 28]. To assess the percentage of PLA puncta co-localizing with the ER, we employed two different cell lines, our overexpression system (tsA-201 cells) and a rat insulinoma cell line (INS-1) that expresses BK and CaV1.3 channels endogenously. In this and following experiments, we analyze one focal plane, in the middle of the cell, to quantify PLA puncta colocalization with ER membrane. Figure 3D shows PLA puncta in the same space as the ER. Comparing the overexpression and endogenous systems, 63 ± 3% versus 50 ± 6% of PLA puncta were localized at the ER (Figure 3E). To rule out the possibility that the colocalization between ensembles and ER is caused by the fact that the ER labeling occupies a large area of the cell, we labeled the ER exit sites with Sec16-GFP and probed for BK and CaV1.3 ensembles with PLA (Figure 3F). We found 28 ± 3% of BK and CaV1.3 ensembles colocalized at ER exit sites in tsA-201 cells, and 11 ± 2% in INS-1 cells (Figure 3G). Although the percentage of colocalization between ensembles and the ER or the ER exit sites cannot be compared or used to understand the dynamics of the ensembles, these results support the conclusion that ensembles are localized in the ER, and suggest that BK and CaV1.3 channels traffic together through the ER and exit together.

Revised version: We next investigated the identity of the intracellular membranes where these PLA puncta were found. A large component of intracellular membranes in the cell is the endoplasmic reticulum (ER), where channels are inserted after translation. To determine whether BK and CaV1.3 associate in the ER (Figure 3A), we combined PLA with immunodetection. As before, PLA probed for BK and CaV1.3 hetero-clusters, while a KDEL-moxGFP label identified the ER (Figure 3B). To avoid disruption of organelle architecture, we used the monomeric mox version of KDEL-GFP, which is optimized to reduce oligomerization in the ER environment [28]. Neither overexpression of KDEL-moxGFP nor fixation altered the ER structure as the tubule width remained around 150 nm for either condition (Figure 3C), a value in agreement with the literature [29, 30]. To assess the percentage of PLA puncta colocalizing with the ER, we employed two different cell lines, our overexpression system (tsA-201 cells) and a rat insulinoma cell line (INS-1) that expresses BK and CaV1.3 channels endogenously. In this and following experiments, we analyze one focal plane, in the middle of the cell, to quantify PLA puncta colocalization with ER membrane. Figure 3D shows PLA puncta in the same space as the ER. Comparing the overexpression and endogenous systems, 63 ± 3% versus 50 ± 6% of total PLA puncta were localized at the ER (Figure 3E). To determine whether the observed colocalization between BK–CaV1.3 hetero-clusters and the ER was not simply due to the extensive spatial coverage of ER labeling, we labeled ER exit sites using Sec16-GFP and probed for hetero-clusters with PLA. This approach enabled us to test whether the hetero-clusters were preferentially localized to ER exit sites, which are specialized trafficking hubs that mediate cargo selection and direct proteins from the ER into the secretory pathway. In contrast to the more expansive ER network, which supports protein synthesis and folding, ER exit sites ensure efficient and selective export of proteins to their target destinations.

Earlier version: BK and CaV1.3 ensembles localize at ER and ER exit sites (ERES).

Revised version: BK and CaV1.3 hetero-clusters localize at ER and ER exit sites (ERES).

Earlier version: A. Diagram of the hypothesis: Ensembles of BK (magenta) and CaV1.3 (cyan) can be found at the ER membrane. B. Representative image of the ER labeled with exogenous GFP in INS-1 cells. Cells were transfected with KDEL-moxGFP. Enlargement of a selected region is shown in the inset. C. Comparison of the ER tubule distance in live and fixed tsA-201 and INS-1 cells. Data points are from n = 23 tsA-201 cells, n = 27 INS-1 cells D. Confocal images of PLA puncta (magenta) and ER (green) in tsA-201 (left) and INS-1 (right) cells. Cells were transfected with BK, CaV1.3, and KDEL-moxGFP. PLA puncta reveal ensembles of BK and CaV1.3. ER is shown in green. E. Comparison of BK and CaV1.3 ensembles found at the ER and relative to all PLA puncta in the cell. Values are given in percentages. F. Representative image of ER exit sites labeled with sec16-GFP. Magnification is shown in the inset. G. Thresholded images of PLA puncta and ER exit sites. Left and right are the same as in D, but cells were transfected with Sec16-GFP instead of KDEL. H. Percentage of BK and CaV1.3 ensembles found at ER exit sites relative to all PLA puncta. Grey bars show the percentage of ensembles found in the ER. Data points are from n = 45 tsA-201 cells for ER; n = 21 tsA-201 cells for ERES; n = 23 INS-1 cells for ER; n = 23 INS-1 cells for ER exit sites. Scale bar is 10 μm and 2 μm in the magnifications.

Revised version: A. Diagram of the hypothesis: hetero-clusters of BK (magenta) and CaV1.3 (cyan) can be found at the ER membrane. B. Representative image of the ER labeled with exogenous GFP in INS-1 cells. Cells were transfected with KDEL-moxGFP. Magnification is shown in the inset. C. Comparison of the ER tubule distance in live and fixed tsA-201 and INS-1 cells. Data points are from n = 23 tsA-201 cells, n = 27 INS-1 cells D. Representative images of PLA puncta and ER. Left: tsA-201 cells were transfected with BK, CaV1.3 and KDEL-moxGFP. Right: INS-1 cells were transfected only with KDEL-moxGFP. Fixed cells were probed for BK-CaV1.3 hetero-clusters (PLA puncta) and GFP. PLA puncta is shown in magenta. ER is shown in green. E. Comparison of BK-CaV1.3 hetero-clusters found at the ER and relative to all PLA puncta in the cell. Values are given in percentages. F. Representative images of PLA puncta and ER exit sites. Left and right are the same as in D, but cells were transfected with Sec16-GFP instead of KDEL. G. Comparison of BK-CaV1.3 hetero-clusters found at ER exit sites relative to all PLA puncta in the cell. Values are given in percentages. Data points are from n = 45 tsA-201 cells for ER, n = 21 tsA-201 cells for ERES, n = 23 INS-1 cells for ER, and n = 23 INS-1 cells for ER exit sites. Scale bar is 10 μm and 2 μm in the magnifications.

Added in the revised version: By quantifying the proportion of BK and CaV1.3 hetero-clusters relative to total channel expression at ER exit sites, we found 28 ± 3% colocalization in tsA-201 cells and 11 ± 2% in INS-1 cells (Figure 3F). While the percentage of colocalization between hetero-clusters and the ER or ER exit sites alone cannot be directly compared to infer trafficking dynamics, these findings reinforce the conclusion that hetero-clusters reside within the ER and suggest that BK and CaV1.3 channels traffic together through the ER and exit in coordination.

Earlier version: BK and CaV1.3 ensembles go through the Golgi

Revised version: BK and CaV1.3 hetero-clusters go through the Golgi

Earlier version: Channels are modified in the Golgi after synthesis (Figure 4A), so we asked whether PLA puncta formed by BK and CaV1.3 ensembles are found in the Golgi. Using the same strategy, we labeled Golgi with Gal-T-mEGFP (Figure 4B) and detected BK and CaV ensembles using PLA. To confirm that the overexpressed Gal-T-mEGFP labels the Golgi specifically without altering its structure, we compared the region detected using the antibody against GFP to the region detected using a primary antibody against the Golgi protein 58K (Figure 4C). 58K is a Golgi specific peripheral protein that localizes at the cytosolic face of Golgi [29]. The overlay shows that the GFP signal labels the same region as the antibody against 58K. When assessing the presence of BK and CaV1.3 clusters in the Golgi, we found that 31 ± 5% of PLA puncta were localized at the Golgi in the overexpression system and 25 ± 4% in the endogenous cell model (Figure 4E).

Revised version: Channels are modified in the Golgi after synthesis (Figure 4A), so we asked whether PLA puncta formed by BK and CaV1.3 hetero-clusters can be found in the Golgi. Using the same strategy, we labeled Golgi with Gal-T-mEGFP (Figure 4B) and detected BK-CaV1.3 hetero-clusters using PLA. To confirm that the overexpressed Gal-T-mEGFP labels the Golgi specifically without altering its structure, we compared the region detected by the antibody against GFP to the region detected by a primary antibody against the Golgi protein 58K (Figure 4C). 58K is a specific peripheral protein that localizes to the cytosolic face of Golgi [31]. The overlay shows that the GFP signal labels the same region as the antibody against 58K. When assessing the presence of BK-CaV1.3 hetero-clusters in the Golgi, we found that 31 ± 5% of PLA puncta were localized at the Golgi in the overexpression system and 25 ± 4% in the endogenous cell model (Figure 4E).

Earlier version: BK and CaV1.3 ensembles go through the Golgi.

Revised version: BK and CaV1.3 hetero-clusters go through the Golgi.

Earlier version: A. Diagram of the hypothesis: PLA puncta detecting ensembles between BK (magenta) and CaV1.3 (cyan) channels can be found at the Golgi membrane. B. Representative confocal image of the Golgi structure with exogenous GFP in INS-1 cells. Cells were transfected with Gal-T-mEGFP. Enlargement is shown in the inset. C. Representative images of fixed cells co-stained with antibodies against Gal-T-mEGFP (left) in green and 58K-Golgi (middle) in red. Overlay (right). D. Representative confocal images of PLA puncta and Golgi. TsA-201 cells were transfected with BK, CaV1.3 and Gal-T-mEGFP (left), and INS-1 cells were transfected only with Gal-T-mEGFP (right). PLA puncta are shown in magenta. Golgi is shown in green. E. Scatter dot plot of percentage of BK and CaV1.3 ensembles found at the Golgi relative to all PLA puncta in tsA-201 and INS1 cells. Data points are from n = 22 tsA-201 cells and n = 19 INS-1 cells. F. Representative image of PLA puncta and Golgi. tsA-201 cells were transfected with BK, CaV1.3 and Gal-T-mEGFP. Left: PLA was done against BK and 58K Golgi (magenta) and Golgi is shown in green. Right: PLA was done against CaV1.3 and 58K Golgi (magenta) and Golgi is shown in green. Scale bars: 10 μm and 2 μm in panels B and C; 10 μm and 1 μm in panel D; 2 μm in panel F.

Revised version: A. Diagram of the hypothesis: PLA puncta detecting hetero-clusters between BK (magenta) and CaV1.3 (cyan) channels can be found at the Golgi membrane. B. Representative image of the Golgi structure with exogenous GFP in INS-1 cells. Cells were transfected with Gal-T-mEGFP. Enlargement is shown in the inset. C. Representative images of fixed cells co-stained with antibodies against Gal-T-mEGFP in green and 58K-Golgi in red D. Representative images of PLA puncta and Golgi. tsA-201 cells were transfected with BK, CaV1.3 and Gal-T-mEGFP (left), and INS-1 cells were transfected only with Gal-T-mEGFP (right). PLA puncta are shown in magenta. Golgi is shown in green. E. Scatter dot plot of percentages of BK-CaV1.3 hetero-clusters found at the Golgi relative to all PLA puncta in tsA-201 and INS1 cells. Data points are from n = 22 tsA-201 cells and n = 19 INS-1 cells. F. Representative image of PLA puncta and Golgi. tsA-201 cells were transfected with BK, CaV1.3, and Gal-T-mEGFP. Left: PLA was done against BK and 58K Golgi (magenta), and Golgi is shown in green. Right: PLA was done against CaV1.3 and 58K Golgi (magenta), and Golgi is shown in green. G. Diagram illustrating our interpretation of percentages of BK-CaV1.3 hetero-clusters found in the cell. This illustration is based on results shown in Figures 1-4. Percentages were modified to represent overlap of fluorescent signals and limited resolution. We also show that hetero-clusters found in the ER exit sites (ERES) are also accounted in the ER. Scale bars: 10 μm and 2 μm in panels B and C; 10 μm and 1 μm in panel D; 2 μm in panel F.

Removed in the revised version: Diagram illustrating our interpretation of the percentage of BK-CaV1.3 ensembles found in the cell.

Removed in the revised version: This illustration is based on experiments shown in figures 1 to 4, but percentages were modified to represent overlap of fluorescent signals and limited resolution. We also show that the ensembles found in the ER exit sites (ERES) are also accounted in the ER.

Earlier version: We performed controls to confirm that the formation of ensembles between the two channels was not coincidental but rather the result of structural coupling. We tested the formation of PLA puncta between BK channels and the Golgi protein 58K. We also tested for CaV1.3 and 58K. We did not find PLA puncta when the proximity between the channels and 58K was probed (Figure 4F), supporting the idea that PLA puncta between BK and CaV1.3 channels found at the Golgi represent specific coupling.

Revised version: We performed controls to confirm that the formation of hetero-clusters between the two channels was not coincidental but rather the result of structural coupling. We tested the formation of PLA puncta between BK channels and the Golgi protein 58K. We also tested for CaV1.3 and 58K. We did not find PLA puncta when the proximity between the channels and 58K was probed (Figure 4F), supporting the idea that PLA puncta between BK and CaV1.3 channels found at the Golgi represent specific coupling. We selected the Golgi as a control because it represents the final stage of protein trafficking, ensuring that hetero-cluster interactions observed at this point reflect specificity maintained throughout earlier trafficking steps, including within the ER.

Earlier version: It is important to clarify that the percentages provided in this work cannot be added up to understand the distribution of ensembles along the biosynthetic pathway of the channels. The percentage in each membrane compartment was compared to the total percentage of ensembles observed in each cell for that particular experiment. Therefore, there is expected overlap between our measurements due to i) optical resolution and ii) the effect of not comparing all the organelles in the same cell. Considering these limitations, we interpreted our results as follows: about one half of the ensembles are found inside the cell. The other half corresponds to ensembles at the plasma membrane. From the half inside the cell, roughly one third of ensembles are found in the Golgi, and another third in ER exit sites. The remaining percentage of ensembles are found in other regions of the ER and membranes that we did not measure, such as vesicles.

Revised version: It is important to clarify that the percentages provided in this work cannot be added up to understand the distribution of hetero-clusters along the biosynthetic pathway of the channels. The percentage in each membrane compartment was compared to the total percentage of hetero-clusters observed in each cell for that particular experiment. Therefore, there is expected overlap between our measurements due to (1) optical resolution and (2) the effect of not comparing all the organelles in the same cell. Considering these limitations, we interpreted our results as follows: about one half of hetero-clusters are found inside the cell. The other half corresponds to hetero-clusters at the plasma membrane. From the half inside the cell, roughly one third of hetero-clusters is found in the Golgi, and another third is found in ER exit sites (Figure 4G). The remaining hetero-clusters are found in other regions of the ER and in membranes that we did not probe, such as vesicles. Finally, a key limitation of this approach is that we cannot quantify the proportion of total BK or CaV1.3 channels engaged in hetero-clusters within each compartment. The PLA method provides proximity-based detection, which reflects relative localization rather than absolute channel abundance within individual organelles.

Earlier version: BK mRNA and CaV1.3 mRNA co-localize

Revised version: BK mRNA and CaV1.3 mRNA colocalize

Earlier version: How is it that BK and CaV1.3 proteins come in close proximity in membranes of the ER, ER exit sites and the Golgi? To explore the origins of the initial association, we hypothesized that the two proteins are translated near each other, which could be detected as the colocalization of their mRNAs (Figure 5A and B). The experiment was designed to detect single mRNA molecules from INS-1 cells in culture. We performed multiplex in-situ hybridization experiments using the RNAscope fluorescence detection kit to be able to image three mRNAs simultaneously in the same cell and acquired the images in a confocal microscope with high resolution. Supplementary Figure 4.1 shows images completely void of fluorescent signal from INS-1 cells probed for bacterial mRNA (negative control, Figure S5.1A) and images of INS-1 cells probed for mammalian mRNAs expected to be in these cells (positive control, Figure S5.1B). The probes against the mRNAs of interest and tested in this work were designed by Advanced Cell Diagnostics. We confirmed the specificity of the probes by performing in situ hybridization against BK mRNA (KCNMA1), CaV1.3 mRNA (CACNA1D), type 2 ryanodine receptors mRNA (RyR2), and NaV1.7 mRNA (SCN9A) in non-transfected cell lines (Figure S5.1C and D).

Revised version: How is it that BK and CaV1.3 proteins come in close proximity in membranes of the ER, ER exit sites and the Golgi? To explore the origins of the initial association, we hypothesized that the two proteins are translated near each other, which could be detected as the colocalization of their mRNAs (Figure 5A and B). The experiment was designed to detect single mRNA molecules from INS-1 cells in culture. We performed multiplex in situ hybridization experiments using an RNAscope fluorescence detection kit to be able to image three mRNAs simultaneously in the same cell and acquired the images in a confocal microscope with high resolution. To rigorously assess the specificity of this potential mRNA-level organization, we used multiple internal controls. GAPDH mRNA, a highly expressed housekeeping gene with no known spatial coordination with channel mRNAs, served as a baseline control for nonspecific colocalization due to transcript abundance. To evaluate whether the spatial proximity between BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) was unique to functionally coupled channels, we also tested for NaV1.7 mRNA (SCN9A), a transmembrane sodium channel expressed in INS-1 cells but not functionally associated with BK. This allowed us to determine whether the observed colocalization reflected a specific biological relationship rather than shared expression context. Finally, to test whether this proximity might extend to other calcium sources relevant to BK activation, we probed the mRNA of ryanodine receptor 2 (RyR2), another Ca2+ channel known to interact structurally with BK channels [32]. Together, these controls were chosen to distinguish specific mRNA colocalization patterns from random spatial proximity, shared subcellular distribution, or gene expression level artifacts. Figure 5-Supplementary-2 shows images completely void of fluorescent signal from INS-1 cells probed for bacterial mRNA (negative control, Figure 5-Supplementary 2A) and shows images of INS-1 cells probed for mammalian mRNAs expected to be in these cells (positive control, Figure 5-Supplementary 2B). The probes against the mRNAs of interest and tested in this work were designed by Advanced Cell Diagnostics. We confirmed the specificity of the probes by performing in situ hybridization against KCNMA1, CACNA1D, RyR2, and SCN9A in non-transfected cell lines (Figure 5-Supplementary 2C and D).

Earlier version: BK mRNA (KCNMA1), CaV1.3 mRNA (CACNA1D) co-localize.

Revised version: BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) colocalize.

A. Diagram of the hypothesis: KCNMA1 and CACNA1D mRNAs are found in close proximity to be translated in the same neighborhood. B. Images of fluorescent puncta from RNA scope experiments showing KCNMA1 mRNA in magenta, CACNA1D mRNA in cyan, and GAPDH mRNA in green. Right, magnification of three ROI. C. Comparison of mRNA density of KCNMA1, CACNA1D, and GAPDH. D. Correlation plot of mRNA abundance of KCNMA1 and CACNA1D per cell. E. Correlation plot of mRNA abundance of KCNMA1 and GAPDH per cell. F. Comparison of colocalization between KCNMA1 mRNA and mRNA from CACNA1D, GAPDH, and scrambled images of CACNA1D. Data points are from n = 67 cells. Scale bar is 10 μm and 1 μm in the magnifications.

Illustration of RNA scope methodology.

Steps to detect an mRNA sequence (red zipper) consisting of (1) hybridization, (2) pre-amplification, (3) and labeling. Double ZZ probes are shown in green. Dye labeling the mRNA is shown in magenta.

RNA probe validation.

A. Representative images of INS-1 cells subjected to mRNA probes against a bacterial gene (DapB from Bacillus subtilis). B. Representative images of INS-1 cells subjected to mRNA probes against three constitutive mammalian genes (ubiquitin C, cyclophilin B, and a polymerase II subunit). C-F show tsA-201 cells that were naive to exogenous DNA but treated with mRNA probes against the genes for (C) BK channels, (D) CaV1.3 channels, (E) Ryanodine Receptors type 2 (RyR2), and (F) NaV1.7 channels. Scale bar is 10 μm.

Earlier version: Probes against GAPDH mRNA were used as a control in the same cells using the multiplex capability of this design. Transcripts were detected at different expression levels. GAPDH and CACNA1D mRNAs were more abundant (134 and 33 mRNA/100 μm2, respectively) than KCNMA1 mRNA (12 mRNA/100 μm2, Figure 5C). Interestingly, the abundance of KCNMA1 transcripts correlated more with the abundance of CACNA1D transcripts than with the abundance of GAPDH, a standard housekeeping gene (Figure 5D and E). Furthermore, KCNMA1 and CACNA1D mRNA colocalized by 60 ± 4%, which was 20% more than with GAPDH mRNA (Figure 4F). To rule out the potential influence of difference in mRNA abundance, we assessed CACNA1D mRNA colocalization with randomized, computer-generated KCNMA1 mRNA signals. We found that colocalization between randomized KCNMA1 mRNA and genuine CACNA1D mRNA was also 20% less than the colocalization of genuine KCNMA1 mRNA and genuine CACNA1D mRNA. These results suggest that some fraction of mRNAs for KCNMA1 and CACNA1D channels are translated nearby, so the channel proteins potentially could be inserted into the same regions of the ER. We suggest that the newly synthesized proteins remain together during trafficking through the ER and the Golgi.

Revised version: Probes against GAPDH mRNA were used as a control in the same cells that we probed for KCNMA1 and CACNA1D using the multiplex capability of this design. Transcripts were detected at different expression levels. GAPDH and CACNA1D mRNAs were more abundant (134 and 33 mRNA/100 μm2, respectively) than KCNMA1 mRNA (12 mRNA/100 μm2, Figure 5C). Interestingly, the abundance of KCNMA1 transcripts correlated more with the abundance of CACNA1D transcripts than with the abundance of GAPDH, though with a modest R2 value (Figure 5D and E). Furthermore, KCNMA1 and CACNA1D mRNA colocalized by 60 ± 4%, which was 20% more than with GAPDH mRNA (Figure 5F). As an additional control and to rule out the potential influence of difference between CACNA1D and KCNMA1 abundance, we assessed CACNA1D colocalization against randomized, computer-generated KCNMA1 mRNA signals, where localization was randomized while maintaining the same overall transcript count. The significantly lower (20%) colocalization observed in scrambled conditions compared to genuine BK-Cav1.3 mRNA interactions confirms that proximity is not an artifact of expression levels but reflects a specific spatial association. These results suggest that some fraction of mRNAs for BK and CaV1.3 channels are translated nearby, so the channel proteins potentially could be inserted into the same regions of the ER. We suggest that the newly synthesized proteins remain together during trafficking through the ER and the Golgi.

Earlier version: As an additional control, we tested the proximity between mRNAs of KCNMA1 and voltage-gated sodium channels SCN9A. This sodium channel is expressed in INS-1 cells; however, we do not expect these two mRNAs to translate nearby, as these proteins do not couple functionally. We detected SCN9A mRNA in the same cells where KCNMA1 mRNA was found (Figure 6A). The colocalization between KCNMA1 mRNA and SCN9A mRNA was only 18 ± 2% (Figure 6C), which was less than what was observed with KCNMA1 mRNA and CACNA1D mRNA (60%, Figure 5F). This result shows that KCNMA1 mRNA tends to localize closer to CACNA1D mRNA than to SCN9A mRNA or GAPDH mRNA, suggesting a preference for translating nearby.

Revised version: Next, we compared the proximity between KCNMA1 and SCN9A. We detected SCN9A in the same cells where KCNMA1 was found (Figure 6A). The colocalization between KCNMA1 and SCN9A was only 18 ± 2% (Figure 6C), which was less than what was observed with KCNMA1 and CACNA1D (60%, Figure 5F). This result shows that KCNMA1 tends to localize closer to CACNA1D than to SCN9A or GAPDH, supporting the specificity of KCNMA1 and CACNA1D mRNA association.

Earlier version: BK (KCNMA1) and RyR-2 (RyR2) mRNA co-localize.

Revised version: BK mRNA (KCNMA1) and RyR-2 mRNA (RyR2) colocalize.

A. Representative confocal images of KCNMA1 and NaV1.7 (SCN9A) mRNA. B. Representative images of KCNMA1 and RyR2 mRNA. C. Comparison of the colocalization between KCNMA1 mRNA and mRNA from RyR2, SCN9A, and scrambled images of KCNMA1. Data points are from n =67 cells. One way ANOVA. Scale bar is 10 μm and 1 μm in the magnifications.

Earlier version: With the goal of understanding if this concept could apply to other channels, we used the same approach to test a second protein known to provide calcium for BK channel opening. Similar to the coupling between BK and CaV1.3 channels, RyR2 also can be structurally coupled to BK channels [29]. Figure 5B shows high-resolution images of single-molecule in-situ hybridization for RyR2 probed together with KCNMA1 in INS-1 cells. In the same population of cells, KCNMA1 mRNA and RyR2 colocalization was 27 ± 3%, which is 1.5 times that with SCN9A mRNA. We also assessed RyR2 colocalization with randomized, computer-generated KCNMA1 mRNA signals. We found that colocalization between randomized KCNMA1 mRNA and genuine RyR2 mRNA was 87% less colocalized compared to the analysis of genuine KCNMA1 mRNA and genuine RyR2 mRNA (Figure 6C), suggesting that KCNMA1 mRNA not only colocalizes with CACNA1D but also with mRNA of other known calcium sources.

Revised version: With the goal of understanding if this concept could apply to other channels, we used the same approach to test a second protein known to provide calcium for BK channel opening. Similar to the coupling between BK and CaV1.3 channels, RyR2 also can be structurally coupled to BK channels [32]. Figure 6B shows high resolution images of single molecule in situ hybridization for RyR2 probed together with KCNMA1 in INS-1 cells. In the same population of cells, KCNMA1 and RyR2 colocalization was 27 ± 3%, which is 1.5 times that with SCN9A mRNA. We also assessed RyR2 colocalization with randomized, computer-generated KCNMA1 mRNA signals. We found that colocalization between randomized KCNMA1 and genuine RyR2 was 87% less colocalized compared to the analysis of genuine KCNMA1 and genuine RyR2 (Figure 6C), suggesting that KCNMA1 not only colocalizes with CACNA1D but also with mRNA of other known calcium sources.

Added in the revised version: To further investigate whether KCNMA1 and CACNA1D are localized in regions of active translation (Figure 7A), we performed RNAscope targeting KCNMA1 and CACNA1D alongside immunostaining for BK protein. This strategy enabled us to visualize transcript-protein colocalization in INS-1 cells with subcellular resolution. By directly evaluating sites of active BK translation, we aimed to determine whether newly synthesized BK protein colocalized with CACNA1D mRNA signals (Figure 7A). Confocal imaging revealed distinct micro-translational complex where KCNMA1 mRNA puncta overlapped with BK protein signals and were located adjacent to CACNA1D mRNA (Figure 7B). Quantitative analysis showed that 71 ± 3% of all KCNMA1 colocalized with BK protein signal which means that they are in active translation. Interestingly, 69 ± 3% of the KCNMA1 in active translation colocalized with CACNA1D (Figure 7C), supporting the existence of functional micro-translational complexes between BK and CaV1.3 channels.

Added in the revised version: BK mRNA (KCNMA1) and CaV1.3 mRNA (CACNA1D) colocalize in micro-translational complexes.

Added in the revised version: A. Diagram of the hypothesis: KCNMA1 mRNAs are found in micro-translational complexes. B. Representative images of KCNMA1 mRNA in magenta, CACNA1D mRNA in cyan, and BK protein in green. C. Comparison of the frequency of colocalization KCNMA1 mRNA in active translation and in micro-translational complexes. Data points are from n = 57 cells. One way ANOVA was used as statistical analysis. Scale bar is 10 μm and 1 μm in the magnifications.

Earlier version: Formation of BK and CaV1.3 ensembles in INS-1 cells

Revised version: BK and CaV1.3 channels form hetero-clusters at the plasma membrane of INS-1 cells

Added in the revised version: We previously showed the organization of BK and CaV1.3 channels in hetero-clusters in tsA-201 cells and neurons, including hippocampal and sympathetic motor neurons [33]. Our present study utilizes INS-1 cells. We detected single molecules from BK and CaV1.3 channels in INS-1 cells using super-resolution microscopy and determined their degree of colocalization. Figure 8A shows a representative image of the localizations of antibodies against BK and CaV1.3 channels.

Removed in the revised version: We previously showed the organization of BK and CaV1.3 channels in ensembles in tsA-201 cells and neurons, including hippocampal and sympathetic motor neurons [16]. Our present study utilized INS-1 cells. We detected single molecules from BK and CaV1.3 channels in INS-1 cells using super-resolution microscopy and determined their degree of co-localization. Figure 7A shows a representative image of the localizations of antibodies against BK and CaV1.3 channels. Maps were rendered at 5 nm, the pixel size for the images presented. BK-positive pixels formed multi-pixel aggregates, which has been interpreted as clusters of BK channels. The same was true for CaV1.3 channels. When looking at the spatial relation between BK and CaV1.3 channels, no fixed geometry or stoichiometry was observed. On the contrary, maps showed a distribution of distances between BK and CaV1.3 detected pixels. In some cases, BK and CaV1.3 channels were perfectly overlapping, as if the detection of the antibodies could not be separated by the 20 nm resolution, suggesting that the same channel aggregate contained BK and CaV1.3 channels. In other cases, BK and CaV1.3 channels were adjacent (less than 5 nm), and in other cases, BK clusters were surrounded by a variable number of CaV1.3 channels.

Earlier version: Formation of BK and CaV1.3 ensembles in INS-1 cells.

Revised version: Formation of BK and CaV1.3 hetero-clusters in INS-1 cells.

Earlier version: (A) Representative localization map of antibodies against BK (magenta) or CaV (cyan) channels. Magnifications are shown in the insets on the right. (B) Scatter dot plot of channel density of BK and CaV1.3 channels in INS-1 cells. (C) Cumulative frequency distribution of the size of BK and CaV1.3 channels. Inset compares median size of BK and CaV channels. (D) Comparison of colocalization between BK and CaV1.3 and BK and scrambled BK images. Data points are from n = 13 cells. Scale bar is 5 μm and 300 nm in the magnifications.

Revised version: A. Representative localization map of antibodies against BK (magenta) and CaV1.3 (cyan) channels. Magnifications are shown in the insets on the right. B. Scatter dot plot of homo-cluster densities of BK and CaV1.3 channels in INS-1 cells. C. Cumulative frequency distributions of homo-cluster sizes of BK and CaV1.3 channels. Inset compares median size of BK and CaV1.3 homo-clusters. D. Comparison of colocalization between BK and CaV1.3 and between scrambled BK and CaV1.3. Data points are from n = 13 cells. Scale bar is 5 μm and 300 nm in the magnifications.

Added in the revised version: Maps were rendered at 5 nm, the pixel size for the images presented. BK-positive pixels formed multi-pixel aggregates, which has been interpreted as homo-clusters of BK channels. CaV1.3 channels also formed homo-clusters. When looking at the formation of BK and CaV1.3 hetero-clusters, no fixed geometry or stoichiometry was observed. On the contrary, maps showed a distribution of distances between BK and CaV1.3 detected pixels. In some cases, BK and CaV1.3 channels were perfectly overlapping, as if the detection of the antibodies could not be separated by the 20 nm resolution. In other cases, BK and CaV1.3 channels were adjacent (less than 5 nm), and in other cases, BK homo-clusters were surrounded by a variable number of CaV1.3 channels.

Earlier version: To provide a more quantitative description of the maps, we measured cluster density, cluster size, and colocalization. We used the particle-analysis tool of the software ImageJ for these measurements. In this analysis, a particle represents positive pixels irrespective of the number of pixels composed. As a note, a particle of one pixel would likely to be a single channel, whereas a particle composed of many pixels would likely to be a cluster of channels. INS-1 cells showed 4 times as many CaV1.3 particles as BK particles (Figure 7B). Figure 6C shows the frequency distribution of particle size, where the median area is 1691 nm2 for BK and 975 nm2 for CaV1.3. We also found 7.5% colocalization between BK and CaV1.3 particles (Figure 7D), which was higher than the colocalization calculated using scrambled images.

Revised version: To provide a more quantitative description of the maps, we measured cluster density, cluster size, and colocalization. We used the particle analysis tool of the software ImageJ for these measurements. In this analysis, a particle for represents positive pixels irrespective of the number of pixels composed. As a note, a particle of one pixel would likely to be a single channel, whereas a particle composed of many pixels would likely to be a cluster of channels. INS-1 cells showed 4 times as many CaV1.3 particles as BK particles (Figure 8B). Figure 8C shows the frequency distribution of particle size, where the median area is 1691 nm2 for BK and 975 nm2 for CaV1.3. We also found 7.5% colocalization between BK and CaV1.3 particles (Figure 8D), which was higher than the colocalization calculated using scrambled images.

Earlier version: Ensembles of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression begins

Revised version: Hetero-clusters of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression

Earlier version: In light of our results, our current model is that BK and CaV1.3 ensembles form prior to their insertion into the plasma membrane. One prediction based on this model is that channels inserted in the plasma membrane would appear as ensembles already at early time points after the start of their synthesis. To test this prediction, we transfected tsA-201 cells with BK and CaV1.3 channels and performed a chase experiment to detect their presence at the plasma membrane using super-resolution microscopy. We measured particle size, density, and colocalization 18, 24, and 48 hours after transfection. Notably, although the channels were transfected simultaneously, BK particles were detected in the plasma membrane 18 hours after transfection, whereas CaV1.3 particles were detected only 24 hours after transfection (Figure 8A-C). The distribution of BK and CaV1.3 particle size did not change with the time after transfection (Figure 8D and E), in agreement with Sato et al., 2019 [24].

Revised version: In light of our results, our current model is that BK and CaV1.3 hetero-clusters form prior to their insertion into the plasma membrane. One prediction based on this model is that channels inserted in the plasma membrane would appear as hetero-clusters already at early time points after the start of their synthesis. To test this prediction, we transfected tsA-201 cells with BK and CaV1.3 channels and performed a chase experiment to detect their presence at the plasma membrane using super-resolution microscopy. We measured particle size, density, and colocalization at 18, 24, and 48 hours after transfection. Notably, although the channels were transfected simultaneously, BK particles were detected in the plasma membrane 18 hours after transfection, whereas CaV1.3 particles were detected 24 hours after transfection (Figure 9A-C). The distribution of BK and CaV1.3 particle size did not change with the time after transfection (Figure 9D and E), in agreement with Sato et al., 2019 [24].

Earlier version: Ensembles of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression begins.

Revised version: Hetero-clusters of BK and CaV1.3 channels are detected at the plasma membrane soon after their expression begins.

Earlier version: (A-C) Representative localization maps of antibodies against BK and CaV channels (A) 18 hours, (B) 24 hours, or (C) 48 hours after DNA transfection into the cell. Enlargements are shown in insets. (D) Size distribution of BK clusters at 18, 24, and 48 hours. Inset compares median BK cluster areas. (E) Size distribution of CaV clusters at 24 and 48-hour time points. Inset compares median CaV cluster areas. CaV1.3 clusters are not present at the 18-hour time point. (F) Comparison of colocalization plots between BK and CaV1.3 channels at 24 and 48-hour time points. Data points are from n = 10 cells. Scale bar is 10 μm and 300 nm in enlargements.

Revised version: A-C. Representative localization maps of antibodies against BK and CaV1.3 channels. (A) 18 hours, (B) 24 hours, or (C) 48 hours after DNA transfection into cells. Enlargements are shown in insets. D. Cumulative frequency distributions of BK homo-cluster size at 18, 24, and 48 hours. Inset compares median BK homo-cluster areas. E. Cumulative frequency distributions of CaV1.3 homo-clusters at 24 and 48 hours. Inset compares median CaV1.3 cluster areas. CaV1.3 clusters are not present at the 18-hour time point. F. Comparison of colocalization plots between BK and CaV1.3 channels at 24 and 48-hour time points. Data points are from n = 10 cells. Scale bar is 10 μm and 300 nm in enlargements.

Earlier version: In support of our prediction, we found particles of CaV1.3 near BK particles as soon as the expression of CaV1.3 channels was detected in the plasma membrane (24 hours after transfection, Enlarged Figure 8B). To test this hypothesis further, we compared plots of localization preference at 24 and 48 hours. These plots were constructed by measuring the percentage of CaV1.3 area occupying concentric regions of 20 nm width around BK particles. Values were normalized to the percentage of colocalization found at 200 nm from BK. Figure 8F shows that the localization preference plots are identical at 24 and 48 hours, consistent with the idea that these channels are inserted together already as ensembles.

Revised version: In support of our prediction, we found particles of CaV1.3 near BK particles as soon as the expression of CaV1.3 channels was detected in the plasma membrane (24 hours after transfection, enlarged Figure 9B). To test this hypothesis further, we compared plots of localization preference at 24 and 48 hours. These plots were constructed by measuring the percentage of CaV1.3 area occupying concentric regions of 20 nm width around BK particles. Values were normalized to the percentage of colocalization found at 200 nm from BK. Figure 9F shows that the localization preference plots are identical at 24 and 48 hours, consistent with the hypothesis that both channels are inserted together into the plasma membrane already as hetero-clusters.

Discussion

Intracellular Assembly of BK and CaV1.3 Channels

Earlier version: Previous work from our group has revealed the organization as ensembles of large-conductance calcium-activated potassium (BK) channels and voltage-gated calcium channels, particularly CaV1.3, within nanodomains. This spatial organization is crucial for the functional coupling that enables BK channels to respond promptly to calcium influx, thereby modulating cellular excitability [15]. While physical details of BK-CaV interactions are not fully elucidated, the evidence suggests that BK channels interact with both the CaV α1 subunit and its auxiliary subunits [30, 31]. Our findings suggest that the co-assembly of BK and CaV1.3 channels into ensembles occurs early in the process of protein synthesis, before arrival at the Golgi. Thus, the channels would be properly configured for immediate functional interaction upon reaching the cell surface.

Revised version: Previous work from our group has revealed the organization of hetero-clusters of large-conductance calcium-activated potassium (BK) channels and voltage-gated calcium channels, particularly CaV1.3, within nanodomains. This spatial organization is crucial for the functional coupling that enables BK channels to respond promptly to calcium influx, thereby modulating cellular excitability [15]. While physical details of BK-CaV1.3 interactions are not fully elucidated, the evidence suggests that BK channels interact with both the CaV α1 subunit and its auxiliary subunits [34, 35]. Our findings highlight the intracellular assembly of BK-CaV1.3 hetero-clusters, though limitations in resolution and organelle-specific analysis prevent precise quantification of the proportion of intracellular complexes that ultimately persist on the cell surface. While our data confirms that hetero-clusters form before reaching the plasma membrane, it remains unclear whether all intracellular hetero-clusters transition intact to the membrane or undergo rearrangement or disassembly upon insertion. Future studies utilizing live cell tracking and high resolution imaging will be valuable in elucidating the fate and stability of these complexes after membrane insertion.

Earlier version: A stochastic model of ion channel cluster formation in the plasma membrane proposes that random yet probabilistic interactions between ion channels contribute to the formation of homo-clusters [24]. These clusters are stabilized and organized within specific membrane regions through biophysical mechanisms such as membrane curvature. Recent findings on spatial organization of G-protein coupled receptors provide additional support for this framework, demonstrating that the coupling of membrane proteins to the curvature of the plasma membrane acts as a driving force for their clustering into functional domains [32]. Yet, the mechanisms regulating the spatial organization of hetero-clusters (ensembles) are less known.

Revised version: A stochastic model of ion channel homo-cluster formation in the plasma membrane proposes that random yet probabilistic interactions between ion channels contribute to their formation [24]. These clusters are stabilized and organized within specific membrane regions through biophysical mechanisms such as membrane curvature. Recent findings on spatial organization of G-protein coupled receptors provide additional support for this framework, demonstrating that the coupling of membrane proteins to the curvature of the plasma membrane acts as a driving force for their clustering into functional domains [36]. Yet, the mechanisms regulating the spatial organization of hetero-clusters are less known.

Added in the revised version: Colocalization and Trafficking Dynamics

Added in the revised version: The colocalization of BK and CaV1.3 channels in the ER and at ER exit sites before reaching the Golgi suggests a coordinated trafficking mechanism that facilitates the formation of multi-channel complexes crucial for calcium signaling and membrane excitability [37, 38]. Given the distinct roles of these compartments, colocalization at the ER and ER exit sites may reflect transient proximity rather than stable interactions. Their presence in the Golgi further suggests that post-translational modifications and additional assembly steps occur before plasma membrane transport, providing further insight into hetero-cluster maturation and sorting events. By examining BK-CaV1.3 hetero-cluster distribution across these trafficking compartments, we ensure that observed colocalization patterns are considered within a broader framework of intracellular transport mechanisms [39]. Previous studies indicate that ER exit sites exhibit variability in cargo retention and sorting efficiency [40], emphasizing the need for careful evaluation of colocalization data. Accounting for these complexities allows for a robust assessment of signaling complexes formation and trafficking pathways.

Added in the revised version: BK Surface Expression and Independent Trafficking Pathways

Added in the revised version: BK surface expression in the absence of CaV1.3 indicates that its trafficking does not strictly rely on CaV1.3-mediated interactions. Since BK channels can be activated by multiple calcium sources, their presence in intracellular compartments suggests that their surface expression is governed by intrinsic trafficking mechanisms rather than direct calcium-dependent regulation. While some BK and CaV1.3 hetero-clusters assemble into signaling complexes intracellularly, other BK channels follow independent trafficking pathways, demonstrating that complex formation is not obligatory for all BK channels. Differences in their transport kinetics further reinforce the idea that their intracellular trafficking is regulated through distinct mechanisms. Studies have shown that BK channels can traffic independently of CaV1.3, relying on alternative calcium sources for activation [13, 41]. Additionally, CaV1.3 exhibits slower synthesis and trafficking kinetics than BK, emphasizing that their intracellular transport may not always be coordinated. These findings suggest that BK and CaV1.3 exhibit both independent and coordinated trafficking behaviors, influencing their spatial organization and functional interactions.

Removed in the revised version: The co-localization of BK and CaV1.3 channels before reaching to Golgi (ER and at ER exit sites) indicates a coordinated trafficking mechanism. This coordination likely facilitates the formation of multi-channel complexes that are essential for precise calcium signaling and membrane excitability. Additionally, the presence of these ensembles in the Golgi suggests that post-translational modifications and further assembly processes occur before the channels are transported to the plasma membrane. The intracellular assembly of BK and CaV1.3 channels underscores the importance in ensuring efficient functional coupling, which is essential for various physiological processes, including the modulation of neurotransmission and control of neuronal firing pattern [33, 34].

Earlier version: mRNA Co-localization and Protein Trafficking

Revised version: mRNA Colocalization and Protein Trafficking

Earlier version: The co-localization of mRNAs encoding BK and CaV1.3 channels suggests a coordinated translation mechanism that facilitates their proximal synthesis and subsequent assembly into functional complexes. As a mechanism, mRNA localization enhances protein enrichment at functional sites, coordinating with translational control to position ion channels at specific subcellular domains, particularly critical in neurons and cardiomyocytes [35, 36]. By synthesizing these channels in close proximity, this process aids their efficient assembly into functional units reaching their target destination which is crucial for the precise targeting and trafficking of ion channels to specific subcellular domains, thereby ensuring proper cellular function and efficient signal transduction [35, 37-39]. Disruption in mRNA localization or protein trafficking can lead to ion channel mislocalization, resulting in altered cellular function and contributing to disease states. In cardiac cells, precise trafficking of ion channels to specific membrane subdomains is crucial for maintaining normal electrical and mechanical functions, and its disruption contributes to heart disease and arrhythmias [40-42] Similarly, in neurons, the localization and translation of mRNA at synaptic sites are essential for synaptic plasticity, and disturbances can lead to neurological disorders [43, 44].

Revised version: The colocalization of mRNAs encoding BK and CaV1.3 channels suggests a coordinated translation mechanism that facilitates their proximal synthesis and subsequent assembly into functional complexes. As a mechanism, mRNA localization enhances protein enrichment at functional sites, coordinating with translational control to position ion channels at specific subcellular domains; ion channel positioning is especially critical for neurons and cardiomyocytes [42, 43]. By synthesizing these channels in close proximity, channels can be efficiently assembled into functional units, making this process crucial for precise targeting and trafficking of ion channels to specific subcellular domains, thereby ensuring proper cellular function and efficient signal transduction [42, 44-46]. Disruption in mRNA localization or protein trafficking can lead to ion channel mislocalization, resulting in altered cellular function and disease [47, 48]. In cardiac cells, precise trafficking of ion channels to specific membrane subdomains is crucial for maintaining normal electrical and mechanical functions, and its disruption contributes to heart disease and arrhythmias [32, 49, 50] Similarly, in neurons, the localization and translation of mRNA at synaptic sites are essential for synaptic plasticity, and disturbances can lead to neurological disorders [51, 52].

Co-translational Regulation and Functional Coordination

Earlier version: The co-translational association of transcripts encoding BK and CaV1.3 ion channels, as well as other known Ca2+ sources like RyR2, may serve as a mechanism to ensure the precise stoichiometry and assembly of functional channel complexes. This precise assembly is also critical for spatially organizing RyR2 and BK channels, as shown in airway myocytes, where the proximity of these channels underlies the efficient activation of BK channels by localized Ca2+ sparks [45]. Co-translational association of transcripts encoding KV1.3 channels was the first example in which the interaction of nascent KV1.3 N-termini facilitates proper tertiary and quaternary structure required for oligomerization [46, 47]. Similarly, co-translational heteromeric association of hERG1a and hERG1b subunits ensures that cardiac IKr currents exhibit the appropriate biophysical properties and magnitude necessary for normal ventricular action potential shaping [39]. The precise mechanism by which BK transcripts are associated with the transcripts of Ca 2+ sources like CaV1.3 and RyR2 channels remain to be elucidated. While it is conceivable that complementary base pairing or tertiary structural interactions play a role, RNA binding proteins (RBPs) are likely key mediators of these associations.

Revised version: The co-translational association of transcripts encoding BK and CaV1.3 ion channels, as well as other known calcium sources like RyR2, may serve as a mechanism to ensure the precise stoichiometry and assembly of functional channel complexes. This precise assembly is also critical for spatially organizing RyR2 and BK channels, as shown in airway myocytes, where the proximity of these channels underlies the efficient activation of BK channels by localized Ca2+ sparks [32]. Co-translational association of transcripts encoding KV1.3 channels was the first example in which the interaction of nascent KV1.3 N-termini facilitates proper tertiary and quaternary structure required for oligomerization [53, 54]. Similarly, co-translational heteromeric association of hERG1a and hERG1b subunits ensures that cardiac IKr currents exhibit the appropriate biophysical properties and magnitude necessary for normal ventricular action potential shaping [45]. The precise mechanism by which BK transcripts are associated with the transcripts of calcium sources like CaV1.3 and RyR2 channels remain to be elucidated. While it is conceivable that complementary base pairing or tertiary structural interactions play a role, RNA binding proteins (RBPs) are likely key mediators of these associations.

Conclusion and Future Directions

Earlier version: This study provides novel insights into the organization of BK and CaV1.3 channels in ensembles, highlighting the importance of their assembly within the ER, ER exit sites, and Golgi. Our findings reveal that these channels not only assemble at the plasma membrane but also form significant signaling complexes at intracellular membranes, supporting the idea of their coordinated synthesis and trafficking. These observations underscore the critical role of spatial organization in maintaining cellular excitability and calcium signaling. While our study advances the understanding of assembly of BK and CaV1.3 complexes, several key questions remain unanswered. What molecular machinery drives this co-localization at the mRNA and protein level? How do disruptions to complex assembly contribute to channelopathies and related diseases? Additionally, a deeper investigation into the role of RNA binding proteins in facilitating transcript association and localized translation is warranted.

Revised version: This study provides novel insights into the organization of BK and CaV1.3 channels in hetero-clusters, emphasizing their assembly within the ER, at ER exit sites, and within the Golgi. Our findings suggest that BK and CaV1.3 channels begin assembling intracellularly before reaching the plasma membrane, shaping their spatial organization, and potentially facilitating functional coupling. While this suggests a coordinated process that may contribute to functional coupling, further investigation is needed to determine the extent to which these hetero-clusters persist upon membrane insertion. While our study advances the understanding of BK and CaV1.3 hetero-cluster assembly, several key questions remain unanswered. What molecular machinery drives this colocalization at the mRNA and protein level? How do disruptions to complex assembly contribute to channelopathies and related diseases? Additionally, a deeper investigation into the role of RNA binding proteins in facilitating transcript association and localized translation is warranted.